## Q3 (4 points)
จงเขียนโปรแกรมที่อ่านไฟล์ studentScore.txt ที่มีรายชื่อนักเรียนและคะแนนสอบดังแสดงในตัวอย่างข้างล่างเก็บเป็น dictionary แล้วแสดงชื่อนักเรียนเรียงจากคะแนนสูงสุดไปต่ำสุด กำหนดให้ ไม่มีใครได้คะแนนเท่ากัน

### ตัวอย่างไฟล์ studentScore1.txt: 

Piti 25 <br>
Mana 32 <br>
Ari 19 <br>
Suda 28 <br>
Somsak 26 <br>

### ตัวอย่างไฟล์ studentScore2.txt: 

Manee 22 <br>
Piti 32 <br>
Ton 35 <br>
Tree 21 <br>
Siri 23 <br>
Mon 20 <br>

### ตัวอย่างการทำงาน: 
Choose your studentScore: **studentScore1.txt** <br>
Mana <br>
Suda <br>
Somsak <br>
Piti <br>
Ari <br>

### ตัวอย่างการทำงาน: 
Choose your studentScore: **studentScore2.txt** <br>
Ton <br>
Piti <br>
Siri <br>
Manee <br>
Tree <br>

<!-- grade_id: name, points: 0, grade: false, solution: true -->
```python
def get_dict(file_name):
    file_name = open(file_name)
    all_score = {}
    for i in file_name:
        i = i.split()
        all_score[i[0]] = int(i[1])
    return all_score

def run():
    input_name = input()
    data = get_dict(input_name)
    score=list(data.values())
    score.sort(reverse=True)
    for i in score:
        for NN,SS_ in data.items():
            if i == SS_:
                print(NN)
    

```

```python
'''ตรวจสอบผลลัพธ์ของโปรแกรมด้วยการรันที่ cell นี้'''
run()
```

```python
'''ตรวจสอบผลการคำนวณ'''
from unittest.mock import patch
from io import StringIO

def test(inputs,expected_output): 
    with patch('builtins.input', side_effect=inputs), patch('sys.stdout', new_callable=StringIO) as mock_stdout:
        run()
        assert mock_stdout.getvalue() == expected_output,mock_stdout.getvalue()       
```

<!-- grade_id: cell-159cd7c85a0c7a54, points: 1, grade: true, solution: false -->
```python
'''ตรวจสอบผลการคำนวณและการแสดงผลลัพธ์ให้เหมือนกับโจทย์ที่กำหนดให้'''
assert get_dict('studentScore1.txt') == {'Piti': 25, 'Mana': 32, 'Ari': 19, 'Suda': 28, 'Somsak': 26},get_dict('studentScore1.txt')
```

<!-- grade_id: exact, points: 3, grade: true, solution: false -->
```python
'''ตรวจสอบผลการคำนวณและการแสดงผลลัพธ์ให้เหมือนกับโจทย์ที่กำหนดให้'''
'''Example 1'''
test(['studentScore1.txt'],'Mana\nSuda\nSomsak\nPiti\nAri\n') 
test(['studentScore2.txt'],'Ton\nPiti\nSiri\nManee\nTree\nMon\n') 
```