## Q3 (3 points)
จงเขียนโปรแกรมเพื่อตัดสินผลการประกวดร้องเพลง โดยให้โปรแกรมอ่านคะแนนที่ผู้เข้าประกวดได้รับจากคณะกรรมการ <br>
(เต็ม 10 คะแนนสำหรับกรรมการแต่ละท่าน) และผู้ชมในห้องส่ง (เต็ม 70 คะแนน) <br>
จากไฟล์ score\<n\>.txt (แต่ละบรรทัดประกอบด้วยรหัสผู้ประกวดร้องเพลง คะแนนจากกรรมการคนที่ 1 คะแนนจากกรรมการคนที่ 2 คะแนนจากกรรมการคนที่ 3 และคะแนนจากผู้ชมในห้องส่ง) <br>

กำหนดให้เขียนฟังก์ชัน `get_dict(file_name)` ที่มี parameter เป็นชื่อ file แล้วสร้างและคืนค่า dictionary จากการอ่านไฟล์ โดยมี
- key เป็นสตริงของรหัสผู้เข้าประกวด 
- value เป็นผลรวมคะแนนทั้งหมด

จงเขียนฟังก์ชัน `find_top3(scores)` ที่มี parameter เป็น dictionary ของผู้เข้าประกวดและคะแนน แล้วคืนค่าเป็น list ของผู้มีคะแนนสูงสุดการประกวด 3 อันดับแรก กำหนดให้ไม่มีผู้ใดได้คะแนนเท่ากัน

### ตัวอย่างไฟล์  score1.txt: 

CU01 8 8 9 65 <br>
CU02 6 6 6 45 <br>
CU03 8 8 7 55 <br>
CU04 6 8 6 57 <br>
CU05 9 9 9 66 <br>

### ตัวอย่างไฟล์ score2.txt: 

CU01 8 8 9 65 <br>
CU02 6 6 6 45 <br>
CU03 8 7 7 55 <br>
CU04 6 8 6 57 <br>
CU05 9 9 9 66 <br>
CU06 8 8 8 62 <br>
CU07 10 8 9 67 <br>

### ตัวอย่างไฟล์ score3.txt: 

CU01 8 8 9 65 <br>
CU02 6 6 6 45 <br>
CU03 8 7 7 55 <br>
CU04 6 8 6 57 <br>
CU05 9 9 9 66 <br>
CU06 8 8 8 62 <br>
CU07 10 8 9 67 <br>
CU08 8 7 9 60 <br>
CU09 7 7 7 58 <br>
CU10 9 8 9 60 <br>

### ตัวอย่างการทำงาน: 

Choose your score: **score1.txt**<br>
The winners are :<br>
['CU05', 'CU01', 'CU03']

### ตัวอย่างการทำงาน: 

Choose your score: **score2.txt**<br>
The winners are :<br>
['CU07', 'CU05', 'CU01']

### ตัวอย่างการทำงาน: 

Choose your score: **score3.txt**<br>
The winners are :<br>
['CU07', 'CU01', 'CU05']

<!-- grade_id: name, points: 0, grade: false, solution: true -->
```python
def get_dict(file_name):
    '''สร้างและคืนค่า dictionary จากการอ่านไฟล์'''
    # YOUR CODE HERE
    file_name = open(file_name)
    file_name = file_name.readlines()
    data_n = {}
    count_n = 0
    for i in range(len(file_name)):
        file_name_ = file_name[i].split()
        for f in range(1,5):
            count_n += int(file_name_[f])
        data_n[(file_name[i])[0:4]] = count_n
        count_n =0 
    return data_n
def find_top3(score):
    '''คืน list ของผู้มีคะแนนสูงสุดการประกวด 3 อันดับแรก'''
    # YOUR CODE HERE
    list_of_name = []
    data = get_dict(score)
    sortsco = list(data.values())
    sortsco.sort(reverse=True)
    for i in sortsco:
        for NN,SS in data.items():
            if i == SS:
                list_of_name.append(NN)
    return list_of_name
def run():
    '''รับชื่อไฟล์คะแนน แล้วแสดง list ของผู้ชนะ 3 ลำดับแรก'''
    # YOUR CODE HERE
    in_ = input()
    data = find_top3(in_)
    print(f"The winners are :\n{data[0:3]}")
```

```python
'''ตรวจสอบการสร้าง dict'''
'''ตัวอย่าง get_dict("ใส่ชื่อไฟล์.txt")'''
get_dict('score2.txt')
```

Output:
```
{'CU01': 90,
 'CU02': 63,
 'CU03': 77,
 'CU04': 77,
 'CU05': 93,
 'CU06': 86,
 'CU07': 94}
```

<!-- grade_id: exact, points: 1, grade: true, solution: false -->
```python
'''ทดสอบฟังก์ชัน get_dict(file_name)'''
'''Example 1'''
assert get_dict('score1.txt') == {'CU01': 90, 'CU02': 63, 'CU03': 78, 'CU04': 77, 'CU05': 93}, get_dict('score1.txt')
```

<!-- grade_id: cell-c8599a43c8e1b8d8, points: 1, grade: true, solution: false -->
```python
'''ทดสอบฟังก์ชัน get_dict(file_name)'''
'''Example 2-3'''
assert get_dict('score2.txt') == {'CU01': 90, 'CU02': 63, 'CU03': 77, 'CU04': 77, 'CU05': 93, 'CU06': 86, 'CU07': 94}, get_dict('score2.txt')
assert get_dict('score3.txt') == {'CU01': 90, 'CU02': 63, 'CU03': 77, 'CU04': 77, 'CU05': 87, 'CU06': 86, 'CU07': 94, 'CU08': 84, 'CU09': 79, 'CU10': 86}, get_dict('score3.txt')
```

```python
'''ทดสอบฟังก์ชัน find_top3()'''
'''จากตัวอย่าง dictionary ผู้มีคะแนนสูงสุด 3 ลำดับแรกคือ E, B และ F ตามลำดับ'''
assert find_top3({'A':80, 'B':90, 'C':70, 'D':50, 'E':95, 'F':86}) == ['E', 'B', 'F'], find_top3({'A':80, 'B':90, 'C':70, 'D':50, 'E':95, 'F':86})

```

```python
'''ตรวจสอบผลลัพธ์ของโปรแกรมด้วยการรันที่ cell นี้'''
run()
```

<!-- grade_id: partial, points: 0, grade: true, solution: false -->
```python
'''ตรวจสอบผลการคำนวณ'''
from unittest.mock import patch
from io import StringIO
import math

def test(inputs,expected_output): 
    with patch('builtins.input', side_effect=inputs), patch('sys.stdout', new_callable=StringIO) as mock_stdout:
        run()
        assert mock_stdout.getvalue() == expected_output,'\n-----Your output-----\n'+str(mock_stdout.getvalue())+'\n-----Expected-----\n'+str(expected_output)

```

<!-- grade_id: cell-318efae9bca2069c, points: 1, grade: true, solution: false -->
```python
'''ตรวจสอบผลการคำนวณและการแสดงผลลัพธ์ให้เหมือนกับโจทย์ที่กำหนดให้'''
'''top 3'''
test(['score1.txt'],"The winners are :\n['CU05', 'CU01', 'CU03']\n")
test(['score2.txt'],"The winners are :\n['CU07', 'CU05', 'CU01']\n")
test(['score3.txt'],"The winners are :\n['CU07', 'CU01', 'CU05']\n")
```