## Q1 (2 points)
จงเขียนฟังก์ชัน `get_dict(file_name)` ที่รับ parameter เป็นชื่อไฟล์ แล้วสร้างและคืน dictionary ซึ่งมี
- key เป็นตัวเลขของรหัสสินค้า และ มี 
- value เป็น dictionary ที่มี 
    - key เป็นสตริงของขนาด (`S`, `M`,  `L` หรือ `XL`) และ 
    - value เป็นตัวเลขจำนวนของสินค้า
    
และแสดงผลลัพธ์เป็นข้อมูลใน dictionary ตามตัวอย่างการทำงาน

### ตัวอย่างการทำงาน: 
get_dict(**'stock1.txt'**) <br>
{4125: {'M': 4}, 5645: {'L': 7}, 7845: {'XL': 3}}

### ตัวอย่างการทำงาน: 
get_dict(**'stock2.txt'**)<br>
{7845: {'XL': 3}, 4125: {'S': 1, 'M': 4}, 5645: {'L': 7}, 9874: {'S': 3}}

### ตัวอย่างการทำงาน: 
get_dict('**stock3.txt**')<br>
{4125: {'S': 1, 'M': 4}, 5645: {'L': 7}, 9874: {'S': 8}, 9875: {'M': 8}, 7845: {'S': 10, 'XL': 3}}

<!-- grade_id: name, points: 0, grade: false, solution: true -->
```python
def get_dict(file_name):
    '''สร้าง dictionary ของรหัสและจำนวนสินค้า คืนค่าเป็น dictionary'''
    # YOUR CODE HERE
    file_name = open(file_name,'r')
    all_data = {}
    for i in file_name:
        a,b,c = i.strip().split()
        if int(a) not in all_data:
            all_data.update({int(a):{b:int(c)}})
        elif int(a) in all_data:
            if b in all_data[int(a)]:
                all_data[int(a)][b] += int(c)
            else:
                all_data[int(a)].update({b:int(c)})
    return all_data


```

```python
'''ตรวจสอบผลลัพธ์ของโปรแกรมด้วยการรันที่ cell นี้'''
'''ตัวอย่าง get_dict("ใส่ชื่อไฟล์.txt")'''
get_dict('stock1.txt')
```

<!-- grade_id: exact, points: 1, grade: true, solution: false -->
```python
'''ตรวจสอบผลการคำนวณและการแสดงผลลัพธ์ให้เหมือนกับโจทย์ที่กำหนดให้'''
'''Test 1'''
assert get_dict('stock1.txt') == {4125: {'M': 4}, 5645: {'L': 7}, 7845: {'XL': 3}}, get_dict('stock1.txt')
assert get_dict('stock2.txt') == {7845: {'XL': 3}, 4125: {'S': 1, 'M': 4}, 5645: {'L': 7}, 9874: {'S': 3}}, get_dict('stock2.txt')
```

<!-- grade_id: cell-c73870cae5a59dfa, points: 1, grade: true, solution: false -->
```python
'''ตรวจสอบผลการคำนวณและการแสดงผลลัพธ์ให้เหมือนกับโจทย์ที่กำหนดให้'''
'''Test  ลับ'''

assert get_dict('stock4.txt') == {4125: {'S': 2, 'M': 8}, 5645: {'L': 14}, 9874: {'S': 16}, 9875: {'M': 16}, 7845: {'S': 20, 'XL': 6}}, get_dict('stock4.txt')
```