## Q3 (3 points)
จงเขียนโปรแกรมเพื่อรับตัวเลขหนึ่งตัวระหว่าง [0,9] จากนั้นใช้ list comprehension เพื่อแสดงตัวเลขระหว่าง 1-100 ซึ่งมีเลขที่รับเข้ามาปรากฏอยู่

### ตัวอย่างการทำงาน:
Enter a nember between [0,9] : **0** <br>
[10, 20, 30, 40, 50, 60, 70, 80, 90, 100] <br>

### ตัวอย่างการทำงาน: 
Enter a number between [0,9]: **3** <br>
[3, 13, 23, 30, 31, 32, 33, 34, 35, 36, 37, 38, 39, 43, 53, 63, 73, 83, 93] <br>

### ตัวอย่างการทำงาน: 
Enter a number between [0,9]: **4** <br>
[4, 14, 24, 34, 40, 41, 42, 43, 44, 45, 46, 47, 48, 49, 54, 64, 74, 84, 94] <br>

### ตัวอย่างการทำงาน: 
Enter a number between [0,9]: **8** <br>
[8, 18, 28, 38, 48, 58, 68, 78, 80, 81, 82, 83, 84, 85, 86, 87, 88, 89, 98] <br>

<!-- grade_id: number, points: 0, grade: false, solution: true -->
```python
def run():
    list_num = []
    input_ = input("[0,9] : ")
    for i in range(1,101):
        if input_ in str(i):
            list_num.append(i)
    print(list_num)
```

```python
'''ตรวจสอบผลลัพธ์ของโปรแกรมด้วยการรันที่ cell นี้'''
run()
```

Output:
```
[0,9] : 3
[3, 13, 23, 30, 31, 32, 33, 34, 35, 36, 37, 38, 39, 43, 53, 63, 73, 83, 93]
```

<!-- grade_id: partial, points: 0, grade: true, solution: false -->
```python
'''ตรวจสอบผลการคำนวณ'''
from unittest.mock import patch
from io import StringIO
import math

def test(inputs,expected_output,exact_output): 
    with patch('builtins.input', side_effect=inputs), patch('sys.stdout', new_callable=StringIO) as mock_stdout:
        run()
        assert mock_stdout.getvalue() == expected_output,mock_stdout.getvalue()
```

<!-- grade_id: cell-0c38a636281d4f0d, points: 1, grade: true, solution: false -->
```python
'''ตรวจสอบผลการคำนวณและการแสดงผลลัพธ์ให้เหมือนกับโจทย์ที่กำหนดให้'''
'''Even number'''
test(['0'],"[10, 20, 30, 40, 50, 60, 70, 80, 90, 100]\n",exact_output=True) 
test(['4'],"[4, 14, 24, 34, 40, 41, 42, 43, 44, 45, 46, 47, 48, 49, 54, 64, 74, 84, 94]\n",exact_output=True) 
test(['8'],"[8, 18, 28, 38, 48, 58, 68, 78, 80, 81, 82, 83, 84, 85, 86, 87, 88, 89, 98]\n",exact_output=True) 
```

<!-- grade_id: exact, points: 1, grade: true, solution: false -->
```python
'''ตรวจสอบผลการคำนวณและการแสดงผลลัพธ์ให้เหมือนกับโจทย์ที่กำหนดให้'''
'''Odd number'''
test(['1'],"[1, 10, 11, 12, 13, 14, 15, 16, 17, 18, 19, 21, 31, 41, 51, 61, 71, 81, 91, 100]\n",exact_output=True) 
test(['3'],"[3, 13, 23, 30, 31, 32, 33, 34, 35, 36, 37, 38, 39, 43, 53, 63, 73, 83, 93]\n",exact_output=True) 
test(['9'],"[9, 19, 29, 39, 49, 59, 69, 79, 89, 90, 91, 92, 93, 94, 95, 96, 97, 98, 99]\n",exact_output=True) 
```

<!-- grade_id: cell-8bd692f819aeb5f4, points: 1, grade: true, solution: false -->
```python
'''เทสลับ'''
```