## Q2 (2 points)
จงเขียนฟังก์ชัน sum_digits ซึ่งรับ string ที่ประกอบไปด้วยตัวอักษรและตัวเลข แล้วคืนค่าผลบวกของตัวเลขทั้งหมดที่ปรากฏอยู่ใน string นั้น

### ตัวอย่างการทำงาน: 

Enter your string: zz1xa2d3 <br>
6

### ตัวอย่างการทำงาน: 

Enter your string: 5432d <br>
14

### ตัวอย่างการทำงาน: 

Enter your string: Hello <br>
0

<!-- grade_id: score, points: 0, grade: false, solution: true -->
```python
def sum_digits(text) :
    # YOUR CODE HERE
    sum_ = 0
    for i in text:
        if i in ["1",'2','3','4','5','6','7','8','9','0']:
            sum_ = int(i)+ sum_
    print(sum_)
    
def run():
    '''main program'''
    # YOUR CODE HERE
    text = input('text : ')
    sum_digits(text)
```

```python
'''ตรวจสอบผลลัพธ์ของโปรแกรมด้วยการรันที่ cell นี้'''
run()
```

Output:
```
text : 1231
เข้า
1
เข้า
2
เข้า
3
เข้า
1
sum_ : 7
```

<!-- grade_id: partial, points: 0, grade: true, solution: false -->
```python
'''ตรวจสอบผลการคำนวณ'''
from unittest.mock import patch
from io import StringIO
import math

def test(inputs,expected_output,exact_output): 
    with patch('builtins.input', side_effect=inputs), patch('sys.stdout', new_callable=StringIO) as mock_stdout:
        run()
        assert mock_stdout.getvalue() == expected_output,mock_stdout.getvalue()
```

<!-- grade_id: exact, points: 1, grade: true, solution: false -->
```python
'''ตรวจสอบผลการคำนวณและการแสดงผลลัพธ์ให้เหมือนกับโจทย์ที่กำหนดให้'''
'''Example'''
test(['zz1xa2d3'],'6\n',exact_output=True) 
test(['5432d'],'14\n',exact_output=True) 
test(['Hello'],'0\n',exact_output=True) 
test(['z๒๓๔1xa2d3'],'6\n',exact_output=True) 
```

<!-- grade_id: cell-0c38a636281d4f0d, points: 1, grade: true, solution: false -->
```python
'''เทสลับ'''
```