## Q1 (5 points)
จงเขียนโปรแกรมที่แสดงตัวเลือกให้ผู้ใช้ 6 ตัวเลือกได้แก่ <br>
1 เพิ่มสินค้าไว้ลำดับท้ายสุดในลิสต์, <br>
2 เปลี่ยนสินค้าในลิสต์, <br>
3 เพิ่มสินค้าตามตำแหน่งที่ระบุในลิสต์, <br>
4 ลบสินค้าในลิสต์, <br>
5 แสดงสินค้าในลิสต์, <br>
6 ออกจากโปรแกรม <br>

จากนั้นวนรับตัวเลือกจากผู้ใช้และทำงานตามตัวเลือก โดยให้สร้างเป็นฟังก์ชันสำหรับแต่ละตัวเลือกและรายละเอียดสำหรับแต่ละตัวเลือกมีดังนี้<br>

**การเพิ่มสินค้าไว้ลำดับท้ายสุดในลิสต์** <br>
def add_item(item_list): <br>
โดย item_list คือลิสต์ที่จะเพิ่มสินค้าลงไป <br>
การทำงาน <br>
- รับข้อมูลชื่อสินค้าที่จะเพิ่มจากผู้ใช้ <br>
- หากสินค้านั้นยังไม่ปรากฏในลิสต์ จะเพิ่มสินค้าลงในลิสต์ แล้วแสดงข้อความว่า Item has been added <br>
- หากสินค้านั้นปรากฏในลิสต์อยู่แล้ว จะไม่เพิ่มสินค้าซ้ำลงไป แต่จะแสดงข้อความว่า Item is already in the list  <br>

**การเปลี่ยนสินค้าในลิสต์**  <br>
def change_item(item_list):
การทำงาน
- รับข้อมูลชื่อสินค้าที่จะเปลี่ยนชื่อจากผู้ใช้ <br>
- หากชื่อสินค้านั้นมีในลิสต์ ให้รับชื่อสินค้าใหม่ แล้วเปลี่ยนชื่อสินค้า จากนั้นแสดงข้อความว่า Item has been changed <br>
- หากชื่อสินค้านั้นไม่มีในลิสต์ ให้แสดงข้อความว่า Item is not in the list <br>

**เพิ่มสินค้าตามตำแหน่งที่ระบุในลิสต์**  <br>
def insert_item(item_list):
การทำงาน
- รับข้อมูลชื่อสินค้าที่จะเพิ่มจากผู้ใช้<br>
- รับตำแหน่งที่ต้องการจะเพิ่ม <br>
- เพิ่มชื่อสินค้าตามตำแหน่งที่ระบุ <br>
- แสดงข้อความว่า Item has been inserted<br>

**ลบสินค้าในลิสต์** <br>
def remove_item(item_list): <br>
การทำงาน <br>
- รับข้อมูลชื่อสินค้าที่จะลบจากผู้ใช้ <br>
- หากชื่อสินค้านั้นมีในลิสต์ ให้ลบชื่อสินค้าออกจากลิสต์ จากนั้นแสดงข้อความว่า Item has been removed <br>
- หากชื่อสินค้านั้นไม่มีในลิสต์ ให้แสดงข้อความว่า This item is not in the list <br>

**แสดงสินค้าในลิสต์** <br>
def show_items(item_list): <br>
การทำงาน <br>
- แสดงสินค้าในลิสต์ แต่หากเริ่มแรกสุดยังไม่มีสินค้าให้แสดงข้อความว่า The list is currently empty <br>
 

### ตัวอย่างการทำงาน def show items and exit: 

What would you like to do? <br>
1: add item <br>
2: change item <br>
3: insert item <br>
4: remove item <br>
5: show items <br>
6: exit <br>
---
Enter a number: **5** <br> 
The list is currently empty <br>
Enter a number: **6** <br>

### ตัวอย่างการทำงาน def add item and exit: 
What would you like to do? <br>
1: add item <br>
2: change item <br>
3: insert item <br>
4: remove item <br>
5: show items <br>
6: exit <br>
---
Enter a number: **1** <br>
Enter item: **Pen** <br>
Item has been added <br> 
Enter a number: **1** <br>
Enter item: **Eraser** <br>
Item has been added <br>
Enter a number: **1** <br>
Enter item: **Pen** <br>
Item is already in the list <br>
Enter a number: **5** <br>
['Pen', 'Eraser'] <br>
Enter a number: **6** <br>

### ตัวอย่างการทำงาน def change item and exit: 
What would you like to do? <br>
1: add item <br>
2: change item <br>
3: insert item <br>
4: remove item <br>
5: show items <br>
6: exit <br>
---
Enter a number: **1** <br>
Enter item: **Pen** <br>
Item has been added <br> 
Enter a number: **1** <br>
Enter item: **Eraser** <br>
Item has been added <br>
Enter a number: **2** <br>
Enter item you want to change: **Rubber** <br>
Item is not in the list <br>
Enter a number: **2** <br>
Enter item you want to change: **Eraser** <br>
Enter new item: **Ruler** <br>
Item has been changed <br>
Enter a number: **5** <br>
['Pen', 'Ruler'] <br>
Enter a number: **6** <br>

### ตัวอย่างการทำงาน def insert item  and exit: 
What would you like to do? <br>
1: add item <br>
2: change item <br>
3: insert item <br>
4: remove item <br>
5: show items <br>
6: exit <br>
---
Enter a number: **1** <br>
Enter item: **Pen** <br>
Item has been added <br> 
Enter a number: **1** <br>
Enter item: **Eraser** <br>
Item has been added <br>
Enter a number: **3** <br>
Enter item: **Pencil** <br>
Enter location that you want to insert: **0** <br>
Item has been inserted <br>
Enter a number: **5** <br>
['Pencil', 'Pen', 'Eraser'] <br>
Enter a number: **6** <br>

### ตัวอย่างการทำงาน def remove item  and exit: 
What would you like to do? <br>
1: add item <br>
2: change item <br>
3: insert item <br>
4: remove item <br>
5: show items <br>
6: exit <br>
---
Enter a number: **1** <br>
Enter item: **Pen** <br>
Item has been added <br> 
Enter a number: **1** <br>
Enter item: **Eraser** <br>
Item has been added <br>
Enter a number: **1** <br>
Enter item: **Pencil** <br>
Item has been added <br>
Enter a number: **4** <br>
Enter item you want to remove: **Sharpener** <br>
This item is not in the list <br>
Enter a number: **4** <br>
Enter item you want to remove: **Pen** <br>
Item has been removed <br>
Enter a number: **5** <br>
['Eraser', 'Pencil'] <br>
Enter a number: **6** <br>

<!-- grade_id: name, points: 0, grade: false, solution: true -->
```python
def add_item(item_list) :
    if item_list not in all_item_list:
        all_item_list.append(item_list)    
        print('Item has been added')
    else:
        print('Item is already in the list')
def change_item(item_list) :
    if item_list in all_item_list:
        n=input('Enter new item: ')
        lo=all_item_list.index(item_list)
        all_item_list[lo]=n
        print('Item has been changed')
    else:
        print('Item is not in the list')
def insert_item(item_list) :
    lo=int(input('Enter location that you want to insert: '))
    all_item_list.insert(lo,item_list)
    print('Item has been inserted')
def remove_item(item_list) :
    if item_list in all_item_list:
        all_item_list.remove(item_list)
        print('Item has been removed')
    else:
        print('This item is not in the list')
def show_item(item_list) :
    if len(all_item_list)==0:
        print('The list is currently empty')
    else:
        print(all_item_list)

def run():
    print('What would you like to do?\n1: add item\n2: change item\n3: insert item\n4: remove item\n5: show items\n6: exit')
    global all_item_list
    all_item_list=[]
    while True:
        choice=input('Enter a number: ')
        if choice=='1':
            item=input('Enter item: ')
            add_item(item)
        if choice=='2':
            item=input('Enter item you want to change: ')
            change_item(item)
        if choice=='3':
            item=input('Enter item: ')
            insert_item(item)
        if choice=='4':
            item=input('Enter itemyou want to remove: ')
            remove_item(item)
        if choice=='5':
            show_item('')
        if choice=='6':
            break
```

```python
'''ตรวจสอบผลลัพธ์ของโปรแกรมด้วยการรันที่ cell นี้'''
run()
```

<!-- grade_id: partial, points: 0, grade: true, solution: false -->
```python
'''ตรวจสอบผลการคำนวณ'''
from unittest.mock import patch
from io import StringIO
import math

def test(inputs,expected_output,exact_output): 
    with patch('builtins.input', side_effect=inputs), patch('sys.stdout', new_callable=StringIO) as mock_stdout:
        run()
        assert mock_stdout.getvalue() == expected_output,mock_stdout.getvalue()
```

<!-- grade_id: exact, points: 1, grade: true, solution: false -->
```python
'''ตรวจสอบผลการคำนวณและการแสดงผลลัพธ์ให้เหมือนกับโจทย์ที่กำหนดให้'''
'''Test function show'''
test(['5','6'],'What would you like to do?\n1: add item\n2: change item\n3: insert item\n\
4: remove item\n5: show items\n6: exit\nThe list is currently empty\n',exact_output=True)
```

<!-- grade_id: cell-c73870cae5a59dfa, points: 1, grade: true, solution: false -->
```python
'''ตรวจสอบผลการคำนวณและการแสดงผลลัพธ์ให้เหมือนกับโจทย์ที่กำหนดให้'''
'''Test function add'''
test(['1','Pen','1','Eraser','1','Pen','5','6'],'What would you like to do?\n1: add item\n2: change item\n3: insert item\n\
4: remove item\n5: show items\n6: exit\nItem has been added\nItem has been added\nItem is already in the list\n[\'Pen\', \'Eraser\']\n',exact_output=True)
```

<!-- grade_id: cell-cbc47179f27fefee, points: 1, grade: true, solution: false -->
```python
'''ตรวจสอบผลการคำนวณและการแสดงผลลัพธ์ให้เหมือนกับโจทย์ที่กำหนดให้'''
'''Test function change'''
test(['1','Pen','1','Eraser','2','Rubber','2','Eraser','Ruler','5','6'],'What would you like to do?\n1: add item\n2: change item\n3: insert item\n\
4: remove item\n5: show items\n6: exit\nItem has been added\nItem has been added\nItem is not in the list\n\
Item has been changed\n[\'Pen\', \'Ruler\']\n',exact_output=True)
```

<!-- grade_id: cell-0ec36ba71a93f61f, points: 1, grade: true, solution: false -->
```python
'''ตรวจสอบผลการคำนวณและการแสดงผลลัพธ์ให้เหมือนกับโจทย์ที่กำหนดให้'''
'''Test function insert'''
test(['1','Pen','1','Eraser','3','Pencil','0','5','6'],'What would you like to do?\n1: add item\n2: change item\n3: insert item\n\
4: remove item\n5: show items\n6: exit\nItem has been added\nItem has been added\nItem has been inserted\n[\'Pencil\', \'Pen\', \'Eraser\']\n',exact_output=True)
```

<!-- grade_id: cell-f59ccd5b163fae70, points: 1, grade: true, solution: false -->
```python
'''ตรวจสอบผลการคำนวณและการแสดงผลลัพธ์ให้เหมือนกับโจทย์ที่กำหนดให้'''
'''Test function remove'''
test(['1','Pen','1','Eraser','1','Pencil','4','Sharpener','4','Pen','5','6'],'What would you like to do?\n1: add item\n2: change item\n3: insert item\n\
4: remove item\n5: show items\n6: exit\nItem has been added\nItem has been added\nItem has been added\nThis item is not in the list\nItem has been removed\n\
[\'Eraser\', \'Pencil\']\n',exact_output=True)
```