## Q1 (2 points)
จงเขียนโปรแกรมที่อ่านเวกเตอร์ 2 เวกเตอร์จากไฟล์ vectors\<n\>.txt แล้วตรวจสอบว่ามีขนาดเท่ากันหรือไม่
- ถ้าไม่เท่ากัน จะแสดงข้อความว่า Incompatible size
- ถ้าเท่ากัน จะหา dot product ของเวกเตอร์ทั้งสอง
และแสดงผลลัพธ์ดังตัวอย่างข้างล่างนี้

### ตัวอย่างการทำงาน: 

ไฟล์ vectors1.txt:  <br>
6 4 9 6 -4 7 <br>
3 7 -3 9 -5 1 <br>
---
Choose your vector file: **vectors1.txt** <br>
v1 = [6.0, 4.0, 9.0, 6.0, -4.0, 7.0] <br>
v2 = [3.0, 7.0, -3.0, 9.0, -5.0, 1.0] <br>
v1*v2 = 100.0

### ตัวอย่างการทำงาน: 

ไฟล์ vectors2.txt: <br>
6 4 9 6 -4 7 <br>
3 7 -3 <br>
---
Choose your vector file: **vectors2.txt** <br>
v1 = [6.0, 4.0, 9.0, 6.0, -4.0, 7.0] <br>
v2 = [3.0, 7.0, -3.0] <br>
Incompatible size

### ตัวอย่างการทำงาน: 

ไฟล์ vectors3.txt: <br>
5.5 4.5 9.5 6 -4 7 <br>
3 7.1 -3.5 5 6 9 <br>

Choose your vector file: **vectors3.txt** <br>
v1 = [5.5, 4.5, 9.5, 6.0, -4.0, 7.0] <br>
v2 = [3.0, 7.1, -3.5, 5.0, 6.0, 9.0] <br>
v1*v2 = 84.2

### ตัวอย่างการทำงาน: 

ไฟล์ vectors4.txt: <br>
5.5 4 -9 6 -4 <br>
3 -7.1 0  <br>

Choose your vector file: **vectors4.txt** <br>
v1 = [5.5, 4.0, -9.0, 6.0, -4.0] <br>
v2 = [3.0, -7.1, 0.0] <br>
Incompatible size <br>

<!-- grade_id: vector, points: 0, grade: false, solution: true -->
```python
def run():
    cho_file = input("ชื่อไฟล์ไรอะ :")
    df = open(cho_file,'r')
    df = df.readlines()
    total = 0 
    print('v1 =',list(float(i) for i in (df[0].split())))
    print('v2 =',list(float(i) for i in (df[1].split())))
    if len(df[0].split()) == len(df[1].split()):
        for i in range(len(df[0].split())):
            suma = float((df[0].split())[i])*float((df[1].split())[i])
            total += suma
        print('v1*v2 =',total) 
    else :
        print('Incompatible size')
```

```python
'''ตรวจสอบผลลัพธ์ของโปรแกรมด้วยการรันที่ cell นี้'''
run()
```

Output:
```
ชื่อไฟล์ไรอะ :vector2.txt
Incompatible size
```

<!-- grade_id: partial, points: 0, grade: true, solution: false -->
```python
'''ตรวจสอบผลการคำนวณ'''
from unittest.mock import patch
from io import StringIO
import math

def test(inputs,expected_output,exact_output): 
    with patch('builtins.input', side_effect=inputs), patch('sys.stdout', new_callable=StringIO) as mock_stdout:
        run()
        assert mock_stdout.getvalue() == expected_output,mock_stdout.getvalue()
```

<!-- grade_id: exact, points: 1, grade: true, solution: false -->
```python
'''ตรวจสอบผลการคำนวณและการแสดงผลลัพธ์ให้เหมือนกับโจทย์ที่กำหนดให้'''
'''Vector dot product'''
test(['vector1.txt'],'v1 = [6.0, 4.0, 9.0, 6.0, -4.0, 7.0]\nv2 = [3.0, 7.0, -3.0, 9.0, -5.0, 1.0]\nv1*v2 = 100.0\n',exact_output=True)
test(['vector3.txt'],'v1 = [5.5, 4.5, 9.5, 6.0, -4.0, 7.0]\nv2 = [3.0, 7.1, -3.5, 5.0, 6.0, 9.0]\nv1*v2 = 84.2\n',exact_output=True)
```

<!-- grade_id: cell-cbc47179f27fefee, points: 1, grade: true, solution: false -->
```python
'''ตรวจสอบผลการคำนวณและการแสดงผลลัพธ์ให้เหมือนกับโจทย์ที่กำหนดให้'''
'''Incompatible size'''
test(['vector2.txt'],'v1 = [6.0, 4.0, 9.0, 6.0, -4.0, 7.0]\nv2 = [3.0, 7.0, -3.0]\nIncompatible size\n',exact_output=True)
test(['vector4.txt'],'v1 = [5.5, 4.0, -9.0, 6.0, -4.0]\nv2 = [3.0, -7.1, 0.0]\nIncompatible size\n',exact_output=True)
```