## Q2 (4 points)
จงเขียนโปรแกรมที่อ่านเฉลยของข้อสอบปรนัยจากไฟล์ sol\<n\>.txt (ข้อสอบมี 10 ข้อ ) 
และอ่านคำตอบของนักเรียนจากไฟล์ exam\<n\>.txt ที่มีคำตอบจากนักเรียนแต่ละคนอยู่ใน 1 บรรทัด
(ถือว่านักเรียนต้องตอบข้อสอบทุกข้อ)
และมีนักเรียนอยู่ 8 คน จากนั้นให้ตรวจคำตอบของนักเรียนทุกคน แล้วหา
- คะแนนของนักเรียนแต่ละคน
- คะแนนเฉลี่ยของข้อสอบแต่ละข้อ
- ข้อสอบข้อที่นักเรียนทำคะแนนได้น้อยสุด (แสดงให้ครบทุกข้อ หากคะแนนเท่ากันหลายข้อ )
- ข้อสอบข้อที่นักเรียนทำคะแนนได้มากสุด (แสดงให้ครบทุกข้อ หากคะแนนเท่ากันหลายข้อ )

**หมายเหตุ ให้ใช้ลิสต์ในข้อนี้ และสามารถใช้ฟังก์ชัน min และ max เพื่อหาค่าน้อยสุดและมากสุดในลิสต์ได้**

### ตัวอย่างไฟล์: 
ไฟล์ sol1.txt: <br>
A A B A D C C D A B<br>
---
ไฟล์ exam1.txt: <br>
A A D A D C C D A B <br>
C A C A D C C D B C <br>
A A B A D C B D A B <br>
B A B A D C C C A B <br>
A A D B D C B D A B <br> 
B A B A D C C C A C <br>
A A D A D C C D A B <br>
A A B A C C C D B B <br>
---
ไฟล์ sol2.txt: <br>
C A B A D A C C A B <br>
---
ไฟล์ exam2.txt: <br>
A A B A C C C D B B <br>
A A B A D C B D A B <br>
B A C A D B C D B C <br>
A A D B D C B D A B <br>
A A B A D C C D A B <br>
A A D A D C C D A B <br>
C A B A D A C C A C <br>
A A D A D C C D A B <br>

### ตัวอย่างการทำงาน: 
Choose your solution file: **sol1.txt** <br>
Choose your exam file: **exam1.txt** <br>

Student score : [9, 6, 9, 8, 7, 7, 9, 8]<br>

Question 1 : 0.625 <br>
Question 2 : 1.0 <br>
Question 3 : 0.5 <br>
Question 4 : 0.875 <br>
Question 5 : 0.875 <br>
Question 6 : 1.0 <br>
Question 7 : 0.75 <br>
Question 8 : 0.75 <br>
Question 9 : 0.75 <br>
Question 10 : 0.75 <br>

Hardest : 3 <br>
Easiest : 2 6 <br>

<!-- grade_id: score, points: 0, grade: false, solution: true -->
```python
def run():
    cho_sol = input("เลือกเฉลยทีขอร้อง: ")
    cho_ex = input('เลือกคำตอบนะ: ')
    cho_sol = open(cho_sol,'r')
    cho_ex = open(cho_ex,'r')
    student_score = []
    question =[]
    minn =[]
    maxx =[]
    cho_sol = cho_sol.readline().split()
    cho_ex = cho_ex.readlines()
    n = 0
    numbera = 0
    for i in range(0,8):
        for j in range(0,10):
            if cho_sol[j] == cho_ex[i].split()[j]:
                n += 1
        student_score.append(n)
        n =0
    for i in range(0,8):
        for j in range(0,10):
            for h in range(0,8):
                if cho_sol[j] == cho_ex[h].split()[j]:
                    n +=1 
            if numbera <10:
                question.append(n/8)
                n=0
                numbera +=1
    print('Student score :',student_score)
    print('')
    for i in range(1,11):
        print('Question',i,':',question[i-1])
    print('')
    n=0
    for i in range(0,10):
        n+=1
        if max(question) == question[i]:
            maxx.append(n)
        elif min(question) == question[i]:
            minn.append(n)
    print('Hardest :',end='')
    for i in minn:
        print('',i,end='')
    print('')
    print("Easiest :",end='')
    for i in maxx:
        print('',i, end='')
    print("")
run()
```

```python
'''ตรวจสอบผลลัพธ์ของโปรแกรมด้วยการรันที่ cell นี้'''
run()
```

<!-- grade_id: partial, points: 0, grade: true, solution: false -->
```python
'''ตรวจสอบผลการคำนวณ'''
from unittest.mock import patch
from io import StringIO
import math

def test(inputs,expected_output,exact_output): 
    with patch('builtins.input', side_effect=inputs), patch('sys.stdout', new_callable=StringIO) as mock_stdout:
        run()
        assert mock_stdout.getvalue() == expected_output,mock_stdout.getvalue()
        
```

<!-- grade_id: exact, points: 1, grade: true, solution: false -->
```python
'''ตรวจสอบผลการคำนวณและการแสดงผลลัพธ์ให้เหมือนกับโจทย์ที่กำหนดให้'''
'''sol1-exam1'''
test(['sol1.txt','exam1.txt'],'Student score : [9, 6, 9, 8, 7, 7, 9, 8]\n\n\
Question 1 : 0.625\nQuestion 2 : 1.0\nQuestion 3 : 0.5\nQuestion 4 : 0.875\n\
Question 5 : 0.875\nQuestion 6 : 1.0\nQuestion 7 : 0.75\nQuestion 8 : 0.75\n\
Question 9 : 0.75\nQuestion 10 : 0.75\n\nHardest : 3\n\
Easiest : 2 6\n',exact_output=True) 
```

<!-- grade_id: cell-0c38a636281d4f0d, points: 1, grade: true, solution: false -->
```python
'''ตรวจสอบผลการคำนวณและการแสดงผลลัพธ์ให้เหมือนกับโจทย์ที่กำหนดให้'''
'''sol1-exam2'''
test(['sol1.txt','exam2.txt'],'Student score : [8, 9, 5, 7, 10, 9, 6, 9]\n\n\
Question 1 : 0.75\nQuestion 2 : 1.0\nQuestion 3 : 0.5\nQuestion 4 : 0.875\n\
Question 5 : 0.875\nQuestion 6 : 0.75\nQuestion 7 : 0.75\n\
Question 8 : 0.875\nQuestion 9 : 0.75\nQuestion 10 : 0.75\n\n\
Hardest : 3\nEasiest : 2\n',exact_output=True) 
```

<!-- grade_id: cell-fd46675f2047b1a8, points: 1, grade: true, solution: false -->
```python
'''ตรวจสอบผลการคำนวณและการแสดงผลลัพธ์ให้เหมือนกับโจทย์ที่กำหนดให้'''
'''sol2-exam1'''
test(['sol2.txt','exam1.txt'],'Student score : [6, 5, 6, 8, 4, 7, 6, 5]\n\n\
Question 1 : 0.125\nQuestion 2 : 1.0\nQuestion 3 : 0.5\nQuestion 4 : 0.875\n\
Question 5 : 0.875\nQuestion 6 : 0.0\nQuestion 7 : 0.75\nQuestion 8 : 0.25\n\
Question 9 : 0.75\nQuestion 10 : 0.75\n\nHardest : 6\n\
Easiest : 2\n',exact_output=True) 
```

<!-- grade_id: cell-bdf63c4f416d44e6, points: 1, grade: true, solution: false -->
```python
'''ตรวจสอบผลการคำนวณและการแสดงผลลัพธ์ให้เหมือนกับโจทย์ที่กำหนดให้'''
'''sol2-exam2'''
test(['sol2.txt','exam2.txt'],'Student score : [5, 6, 4, 4, 7, 6, 9, 6]\n\n\
Question 1 : 0.125\nQuestion 2 : 1.0\nQuestion 3 : 0.5\nQuestion 4 : 0.875\n\
Question 5 : 0.875\nQuestion 6 : 0.125\nQuestion 7 : 0.75\nQuestion 8 : 0.125\n\
Question 9 : 0.75\nQuestion 10 : 0.75\n\nHardest : 1 6 8\n\
Easiest : 2\n',exact_output=True) 
```