## Q1 (4 points)
จงเขียนโปรแกรมที่รับตัวเลือก ( circle, square, rectangle ) จากผู้ใช้ แล้วรับความยาวรัศมีของ
วงกลม หรือ ความยาวด้านของสี่เหลี่ยมจัตุรัส หรือ สี่เหลี่ยมผืนผ้า แล้วคำนวณหาพื้นที่และแสดงผลลัพธ์ดัง
ตัวอย่างข้างล่างนี้ <br>
(กำหนดให้ใช้ค่า pi จาก module math)

### ตัวอย่างการทำงาน: 

Choose circle, square or rectangle : **circle** <br>
Length of the radius of the circle : **3.6** <br>
Area is 40.71504079052372

### ตัวอย่างการทำงาน: 

Choose circle, square or rectangle : **square** <br>
Length of the side of the square : **3.6** <br>
Area is 12.96

### ตัวอย่างการทำงาน: 

Choose circle, square or rectangle : **rectangle** <br>
Length of two sides of the rectangle : **2.5, 3.6** <br>
Area is 9.0

### ตัวอย่างการทำงาน: 

Choose circle, square or rectangle : **xx** <br>
Invalid choice.

<!-- grade_id: quardatic, points: 0, grade: false, solution: true -->
```python
import math
def run():
    Ch = input("ch : ")
    if Ch == "circle" :
        x =float(input())
        x = math.pi * (x**2)
        print("Area is "+str(x))
    elif Ch == "square":
        x =float(input())
        x = x**2
        print("Area is "+str(x))
    elif Ch == "rectangle":
        x,y =(float(u) for u in(input().split(",")))
        print("Area is "+str(x*y))
    else :
        print("Invalid choice.")
```

Output:
```
Input In [30]
    print("Area is "x)
                    ^
SyntaxError: invalid syntax
```

```python
'''ตรวจสอบผลลัพธ์ของโปรแกรมด้วยการรันที่ cell นี้'''
run()
```

Output:
```
ch : rectangle
2.5,3.6
9.0
```

<!-- grade_id: partial, points: 0, grade: true, solution: false -->
```python
'''ตรวจสอบผลการคำนวณ'''
from unittest.mock import patch
from io import StringIO
import math

def isfloat(num):
    try:
        float(num)
        return True
    except ValueError:
        return False
    
def test(inputs,expected_output,exact_output): 
    with patch('builtins.input', side_effect=inputs), patch('sys.stdout', new_callable=StringIO) as mock_stdout:
        run()
        assert expected_output==mock_stdout.getvalue(),mock_stdout.getvalue()
                
```

<!-- grade_id: exact, points: 1, grade: true, solution: false -->
```python
'''ตรวจสอบผลการคำนวณและการแสดงผลลัพธ์ให้เหมือนกับโจทย์ที่กำหนดให้'''
test(['circle','3.6'],'Area is 40.71504079052372\n',exact_output=True)   
test(['circle','3.578'],'Area is 40.21893304503943\n',exact_output=True) 
test(['circle','5'],'Area is 78.53981633974483\n',exact_output=True)     
test(['circle','9'],'Area is 254.46900494077323\n',exact_output=True)   
```

<!-- grade_id: cell-40d500747966a4b4, points: 1, grade: true, solution: false -->
```python
'''ตรวจสอบผลการคำนวณและการแสดงผลลัพธ์ให้เหมือนกับโจทย์ที่กำหนดให้'''
test(['square','3.6'],'Area is 12.96\n',exact_output=True) 
test(['square','3.578'],'Area is 12.802083999999999\n',exact_output=True)  
test(['square','5'],'Area is 25.0\n',exact_output=True)  
test(['square','9'],'Area is 81.0\n',exact_output=True)  
```

<!-- grade_id: cell-8ed24f7a3c6c9b8b, points: 1, grade: true, solution: false -->
```python
'''ตรวจสอบผลการคำนวณและการแสดงผลลัพธ์ให้เหมือนกับโจทย์ที่กำหนดให้'''
test(['rectangle','2.5, 3.6'],'Area is 9.0\n',exact_output=True)   
test(['rectangle','3, 4'],'Area is 12.0\n',exact_output=True)      
test(['rectangle','2.5, 4'],'Area is 10.0\n',exact_output=True)    
test(['rectangle','5, 3.6'],'Area is 18.0\n',exact_output=True)    
```

<!-- grade_id: cell-5ea63f69411d8cea, points: 1, grade: true, solution: false -->
```python
'''ตรวจสอบผลการคำนวณและการแสดงผลลัพธ์ให้เหมือนกับโจทย์ที่กำหนดให้'''
test(['xx',' '],'Invalid choice.\n',exact_output=True)
test(['test',' '],'Invalid choice.\n',exact_output=True)
test(['cir',' '],'Invalid choice.\n',exact_output=True)
```