## Q2 (5 points)
จงเขียนโปรแกรมที่รับน้ำหนัก และส่วนสูง จากผู้ใช้ มาคำนวณ bmi โดยให้ระบุหน่วยของน้ำหนักเป็น kg
หรือ lbs และ หน่วยของส่วนสูงเป็น cm หรือ m หรือ ft แล้วไปแปลงเป็น kg และ m และคำนวณแสดงผลลัพธ์ว่าผู้ใช้มีรูปร่างในระดับใด <br>
(1 kg. = 2.205 lbs, 1 m. = 3.2808399 feet)

### ตัวอย่างการทำงาน: 

Weight : **61.3 kg**<br>
Height : **1.65 m**<br>
รูปร่างปกติ

### ตัวอย่างการทำงาน: 

Weight : **159.3 lbs** <br>
Height : **4.92 ft**<br>
อ้วนระดับ 2

### ตัวอย่างการทำงาน:

Weight : **110.8 lbs** <br> Height : **167 cm**<br>ผอม


<!-- grade_id: bmi, points: 0, grade: false, solution: true -->
```python
def run():
    W,T = input().split()
    H,E = input().split()
    a=float(W)
    b=float(H)
    if T == "lbs":
        a = a/2.205
    if E == "ft":
        b = b/3.2808399
    elif E == "cm":
        b = b/100
    c= a/ (b**2)
    if c < 18.5  :
        print("ผอม")
    elif c >= 18.5 and c <23:
        print("รูปร่างปกติ")
    elif c>= 23 and  c<25:
        print("รูปร่างอ้วน")
    elif c>=25 and c< 30:
        print("อ้วนระดับ 1")
    else :
        print("อ้วนระดับ 2")
    
        
```

```python
'''ตรวจสอบผลลัพธ์ของโปรแกรมด้วยการรันที่ cell นี้'''
run()
```

Output:
```
91.5 lbs
128 cm
1.28
41.49659863945578
25.32751381802721
อ้วนระดับ 1
```

<!-- grade_id: partial, points: 0, grade: true, solution: false -->
```python
'''ตรวจสอบผลการคำนวณ'''
from unittest.mock import patch
from io import StringIO
import math

def test(inputs,expected_output,exact_output): 
    with patch('builtins.input', side_effect=inputs), patch('sys.stdout', new_callable=StringIO) as mock_stdout:
        run()
        if exact_output:
            assert mock_stdout.getvalue() == expected_output,mock_stdout.getvalue()
            
```

<!-- grade_id: exact, points: 1, grade: true, solution: false -->
```python
'''ตรวจสอบผลการคำนวณและการแสดงผลลัพธ์ให้เหมือนกับโจทย์ที่กำหนดให้'''
test(['61.3 kg','1.65 m'],'รูปร่างปกติ\n',exact_output=True)  
test(['61.3 kg','165 cm'],'รูปร่างปกติ\n',exact_output=True) 
test(['49.5 kg','5 ft'],'รูปร่างปกติ\n',exact_output=True) 
test(['110.8 lbs','150 cm'],'รูปร่างปกติ\n',exact_output=True) 
test(['110.8 lbs','1.50 m'],'รูปร่างปกติ\n',exact_output=True) 
test(['92.5 lbs','4.8 ft'],'รูปร่างปกติ\n',exact_output=True) 
```

<!-- grade_id: cell-7074c62b5ac8ac0c, points: 1, grade: true, solution: false -->
```python
'''ตรวจสอบผลการคำนวณและการแสดงผลลัพธ์ให้เหมือนกับโจทย์ที่กำหนดให้'''
test(['94 kg','167.64 cm'],'อ้วนระดับ 2\n',exact_output=True)
test(['94.2 kg','1.67 m'],'อ้วนระดับ 2\n',exact_output=True)
test(['61.3 kg','1.5 ft'],'อ้วนระดับ 2\n',exact_output=True)
test(['159.4215 lbs','155 cm'],'อ้วนระดับ 2\n',exact_output=True)
test(['159.4215 lbs','1.5 m'],'อ้วนระดับ 2\n',exact_output=True)
test(['159.3 lbs','4.92 ft'],'อ้วนระดับ 2\n',exact_output=True) 
```

<!-- grade_id: cell-edf90ee888ed0103, points: 1, grade: true, solution: false -->
```python
'''ตรวจสอบผลการคำนวณและการแสดงผลลัพธ์ให้เหมือนกับโจทย์ที่กำหนดให้'''
test(['51.5 kg','167.0 cm'],'ผอม\n',exact_output=True) 
test(['51.3 kg','1.67 m'],'ผอม\n',exact_output=True) 
test(['50.3 kg','5.479002633 ft'],'ผอม\n',exact_output=True) 
test(['110.8 lbs','167 cm'],'ผอม\n',exact_output=True) 
test(['110.8 lbs','1.67 m'],'ผอม\n',exact_output=True) 
test(['24 lbs','5 ft'],'ผอม\n',exact_output=True)
```

<!-- grade_id: cell-401eb542bb74a444, points: 1, grade: true, solution: false -->
```python
'''ตรวจสอบผลการคำนวณและการแสดงผลลัพธ์ให้เหมือนกับโจทย์ที่กำหนดให้'''
test(['92 kg','198 cm'],'รูปร่างอ้วน\n',exact_output=True)
test(['92 kg','2.0 m'],'รูปร่างอ้วน\n',exact_output=True)
test(['92 kg','6.56 ft'],'รูปร่างอ้วน\n',exact_output=True)
test(['150 lbs','165 cm'],'รูปร่างอ้วน\n',exact_output=True)
test(['150 lbs','1.65 m'],'รูปร่างอ้วน\n',exact_output=True)
test(['120 lbs','5 ft'],'รูปร่างอ้วน\n',exact_output=True)
```

<!-- grade_id: cell-0497fd880a4d8063, points: 1, grade: true, solution: false -->
```python
'''ตรวจสอบผลการคำนวณและการแสดงผลลัพธ์ให้เหมือนกับโจทย์ที่กำหนดให้'''
test(['100 kg','2.0 m'],'อ้วนระดับ 1\n',exact_output=True)
test(['100 kg','198 cm'],'อ้วนระดับ 1\n',exact_output=True)
test(['100 kg','6.49 ft'],'อ้วนระดับ 1\n',exact_output=True)
test(['91.5 lbs','128 cm'],'อ้วนระดับ 1\n',exact_output=True)
test(['91.5 lbs','1.28 m'],'อ้วนระดับ 1\n',exact_output=True)
test(['92.5 lbs','4.2 ft'],'อ้วนระดับ 1\n',exact_output=True)
```