## Q2 (3 points)
จงเขียนโปรแกรมที่วนรับชื่อผู้ใช้จนกว่าจะเป็น Lisa หรือใส่ชื่อผิดครบ 3 ครั้ง

### ตัวอย่างการทำงาน: 

Enter username : **Lisa** <br>
Hello, Lisa

### ตัวอย่างการทำงาน: 

Enter username : **test** <br>
Incorrect. Enter again : **Lisa** <br>
Hello, Lisa

### ตัวอย่างการทำงาน:

Enter username : **test** <br>
Incorrect. Enter again : **12345**  <br>
Incorrect. Enter again : **Lisa** <br>
Hello, Lisa

### ตัวอย่างการทำงาน: 

Enter username : **test**  <br>
Incorrect. Enter again : **12345**  <br>
Incorrect. Enter again : **my name** <br>
Not allowed. Incorrect name.

```python

```

<!-- grade_id: name_limit, points: 0, grade: false, solution: true -->
```python
def run():
    s = input("Enter username :")
    i =0 
    while True:
        if s != "Lisa" and i != 2:
                s =input("incorrect, Enter again : ")
                i +=1
        elif i==2 and s!="Lisa":
                print("Not allowed. Incorrect name.")
                break
        else:
            print("Hello,",s) 
            break
```

```python
'''ตรวจสอบผลลัพธ์ของโปรแกรมด้วยการรันที่ cell นี้'''
run()
```

Output:
```
Enter username :'
incorrect, Enter again : '
incorrect, Enter again : '
Not allowed. Incorrect name.
```

<!-- grade_id: partial, points: 0, grade: true, solution: false -->
```python
'''ตรวจสอบผลการคำนวณ'''
from unittest.mock import patch
from io import StringIO
import math

def test(inputs,expected_output,exact_output): 
    with patch('builtins.input', side_effect=inputs), patch('sys.stdout', new_callable=StringIO) as mock_stdout:
        run()
        assert mock_stdout.getvalue() == expected_output,mock_stdout.getvalue()
```

<!-- grade_id: exact, points: 1, grade: true, solution: false -->
```python
'''ตรวจสอบผลการคำนวณและการแสดงผลลัพธ์ให้เหมือนกับโจทย์ที่กำหนดให้'''
test(['Lisa'],'Hello, Lisa\n',exact_output=True) 
test(['.','Lisa'],'Hello, Lisa\n',exact_output=True) 
test(['','Lisa'],'Hello, Lisa\n',exact_output=True) 
test(['test','Lisa'],'Hello, Lisa\n',exact_output=True) 
test(['Blackpink','Lisa'],'Hello, Lisa\n',exact_output=True)  
```

<!-- grade_id: cell-ad91ccd55eeb0c2f, points: 1, grade: true, solution: false -->
```python
'''ตรวจสอบผลการคำนวณและการแสดงผลลัพธ์ให้เหมือนกับโจทย์ที่กำหนดให้'''
test(['test','12345','Lisa'],'Hello, Lisa\n',exact_output=True)
test(['Rose','Jennie','Lisa'],'Hello, Lisa\n',exact_output=True)
test(['Tee','Pee','Lisa'],'Hello, Lisa\n',exact_output=True)
test(['Bambam','Mario','Lisa'],'Hello, Lisa\n',exact_output=True)
```

<!-- grade_id: cell-a07ef84aa3313800, points: 1, grade: true, solution: false -->
```python
'''ตรวจสอบผลการคำนวณและการแสดงผลลัพธ์ให้เหมือนกับโจทย์ที่กำหนดให้'''
test(['test','12345','my name'],'Not allowed. Incorrect name.\n',exact_output=True)
test(['Bambam','Mario','Jennie'],'Not allowed. Incorrect name.\n',exact_output=True)
test(['xxx','56789','Jennie'],'Not allowed. Incorrect name.\n',exact_output=True)
```