## Q4 (5 points)
จงเขียนโปรแกรมที่รับจำนวนนักเรียนที่เข้าสอบ (เป็นจำนวนบวกเท่านั้น) แล้ววนรับคะแนนของนักเรียนที่เข้าสอบให้ครบและหา 
- คะแนนเฉลี่ยของนักเรียนทั้งหมด
- คะแนนเฉลี่ยของนักเรียนที่สอบผ่าน (ได้คะแนนอย่างน้อย 5)
- คะแนนเฉลี่ยของนักเรียนทั้งหมด (ได้คะแนนต่ำกว่า 5)
- คะแนนสูงสุดของนักเรียนทั้งหมด

### ตัวอย่างการทำงาน: 

Number of students : **3** <br>
Student 1 : **10** <br>
Student 2 : **6.5** <br>
Student 3 : **0** <br>
Average score : 5.5 <br>
Average passing score : 8.25 <br>
Average failing score : 0.0 <br>
Highest score : 10.0 <br>

### ตัวอย่างการทำงาน: 

Number of students : **5** <br>
Student 1 : **10** <br>
Student 2 : **6.5** <br>
Student 3 : **0** <br>
Student 4 : **9** <br>
Student 5 : **8.25** <br>
Average score : 6.75 <br>
Average passing score : 8.4375<br>
Average failing score : 0.0 <br>
Highest score : 10.0

### ตัวอย่างการทำงาน: 

Number of students : **7** <br>
Student 1 : **10** <br>
Student 2 : **6.5** <br>
Student 3 : **0** <br>
Student 4 : **9** <br>
Student 5 : **8.25** <br>
Student 6 : **9.50** <br>
Student 7 : **7.25** <br>
Average score : 7.214285714285714 <br>
Average passing score : 8.416666666666666 <br>
Average failing score : 0.0 <br>
Highest score : 10.0 <br>

### ตัวอย่างการทำงาน: 

Number of students : **-1** <br>
Incorrect Input. Number of students : **7** <br>
Student 1 : **10** <br>
Student 2 : **6.5** <br>
Student 3 : **0** <br>
Student 4 : **9** <br>
Student 5 : **8.25** <br>
Student 6 : **9.50** <br>
Student 7 : **7.25** <br>
Average score : 7.214285714285714 <br>
Average passing score : 8.416666666666666 <br>
Average failing score : 0.0 <br>
Highest score : 10.0 <br>

<!-- grade_id: name, points: 0, grade: false, solution: true -->
```python
def run():
    # YOUR CODE HERE
    s = input("Enter :")
    i =0 
    suma = 0
    sumb = 0
    sumc = 0
    maxx = 0
    a = 0
    b = 0
    while True:
        if int(s) < 0 :
            s =float(input("Incorrect Input. Number of students:"))
        if int(s)>0:
            break
    while int(s)>0:
        if i <= int(s)-1 :
            i+=1
            score = float(input("Student :"))
            suma += score
        if maxx <= float(score):
            maxx = score
        if score >= 5 :
            sumb += score
            a += 1 
        if score < 5:
            sumc += score
            b += 1
        if i==int(s):
            break
    print("Average score :",float(suma)/float(s))
    if a ==0:
        print("Average passing score :",sumb)
    elif a>0:
        print("Average passing score :",float(sumb)/int(a))
    if b==0:
        print("Average failing score :",sumc)
    elif b>0:
        print("Average failing score :",float(sumc)/float(b))
    print("Highest score :",float(maxx))
                     
```

```python
'''ตรวจสอบผลลัพธ์ของโปรแกรมด้วยการรันที่ cell นี้'''
run()
```

Output:
```
Enter :1
Student :1
Average score : 1.0
Average passing score : 0
Average failing score : 1.0
Highest score : 1.0
```

<!-- grade_id: partial, points: 0, grade: true, solution: false -->
```python
'''ตรวจสอบผลการคำนวณ'''
from unittest.mock import patch
from io import StringIO
import math

def test(inputs,expected_output,exact_output): 
    with patch('builtins.input', side_effect=inputs), patch('sys.stdout', new_callable=StringIO) as mock_stdout:
        run()
        if exact_output:
            assert mock_stdout.getvalue() == expected_output,mock_stdout.getvalue()
        else:
            assert expected_output in mock_stdout.getvalue(),mock_stdout.getvalue()
```

<!-- grade_id: cell-5646d4ae088951fe, points: 1, grade: true, solution: false -->
```python
'''ตรวจสอบผลการคำนวณและการแสดงผลลัพธ์ให้เหมือนกับโจทย์ที่กำหนดให้'''
'''Number of students is 3'''
test(['3','10','6.5','0'],'Average score : 5.5\nAverage passing score : 8.25\nAverage failing score : 0.0\nHighest score : 10.0\n',exact_output=True) 
test(['3','7','4','9'],'Average score : 6.666666666666667\nAverage passing score : 8.0\nAverage failing score : 4.0\nHighest score : 9.0\n',exact_output=True) 
test(['3','6','5.5','9.8'],'Average score : 7.1000000000000005\nAverage passing score : 7.1000000000000005\nAverage failing score : 0\nHighest score : 9.8\n',exact_output=True) 
test(['3','1.5','4.5','6.5'],'Average score : 4.166666666666667\nAverage passing score : 6.5\nAverage failing score : 3.0\nHighest score : 6.5\n',exact_output=True) 
```

<!-- grade_id: exact, points: 1, grade: true, solution: false -->
```python
'''ตรวจสอบผลการคำนวณและการแสดงผลลัพธ์ให้เหมือนกับโจทย์ที่กำหนดให้'''
'''Number of students is 5'''
test(['5','10','6.5','0','9','8.25'],'Average score : 6.75\nAverage passing score : 8.4375\nAverage failing score : 0.0\nHighest score : 10.0\n',exact_output=True) 
test(['5','6','8','9','4','10'],'Average score : 7.4\nAverage passing score : 8.25\nAverage failing score : 4.0\nHighest score : 10.0\n',exact_output=True) 
test(['5','4.25','6.25','1.25','8.75','9.25'],'Average score : 5.95\nAverage passing score : 8.083333333333334\nAverage failing score : 2.75\nHighest score : 9.25\n',exact_output=True) 
test(['5','4.25','1','0','8.75','9.25'],'Average score : 4.65\nAverage passing score : 9.0\nAverage failing score : 1.75\nHighest score : 9.25\n',exact_output=True) 
```

<!-- grade_id: cell-cbc47179f27fefee, points: 1, grade: true, solution: false -->
```python
'''ตรวจสอบผลการคำนวณและการแสดงผลลัพธ์ให้เหมือนกับโจทย์ที่กำหนดให้'''
'''Number of students is 7'''
test(['7','10','6.5','0','9','8.25','9.50','7.25'],'Average score : 7.214285714285714\nAverage passing score : 8.416666666666666\nAverage failing score : 0.0\nHighest score : 10.0\n',exact_output=True) 
test(['7','4.25','6.25','1.25','8.75','9.25','7.25','9.75'],'Average score : 6.678571428571429\nAverage passing score : 8.25\nAverage failing score : 2.75\nHighest score : 9.75\n',exact_output=True) 
test(['7','6','8','9','4','10','3','4'],'Average score : 6.285714285714286\nAverage passing score : 8.25\nAverage failing score : 3.6666666666666665\nHighest score : 10.0\n',exact_output=True) 
```

<!-- grade_id: cell-f6917d088df9d576, points: 1, grade: true, solution: false -->
```python
'''ตรวจสอบเคส inputs ติด ลบ'''
test(['-1','7','10','6.5','0','9','8.25','9.50','7.25'],'Average score : 7.214285714285714\nAverage passing score : 8.416666666666666\nAverage failing score : 0.0\nHighest score : 10.0\n',exact_output=True) 
test(['-7','-1','7','4.25','6.25','1.25','8.75','9.25','7.25','9.75'],'Average score : 6.678571428571429\nAverage passing score : 8.25\nAverage failing score : 2.75\nHighest score : 9.75\n',exact_output=True) 
```

<!-- grade_id: cell-eb75a0ec803a867e, points: 1, grade: true, solution: false -->
```python
'''เทส ลับ'''
```