## Q3 (7 points)
จงเขียนโปรแกรมที่รับจำนวนนักเรียนที่เข้าสอบ แล้ววนรับคะแนนของนักเรียนที่เข้าสอบให้ครบและจงหาผลรวมคะแนน 

### ตัวอย่างการทำงาน: 

Number of students : **3** <br>
Student 1 : **10** <br>
Student 2 : **6.5** <br>
Student 3 : **0** <br>
Total score : 16.5 <br>

### ตัวอย่างการทำงาน: 

Number of students : **5** <br>
Student 1 : **10** <br>
Student 2 : **6.5** <br>
Student 3 : **0** <br>
Student 4 : **9** <br>
Student 5 : **8.25** <br>
Total score : 33.75 <br>

### ตัวอย่างการทำงาน: 

Number of students : **7** <br>
Student 1 : **10** <br>
Student 2 : **6.5** <br>
Student 3 : **0** <br>
Student 4 : **9** <br>
Student 5 : **8.25** <br>
Student 6 : **9.50** <br>
Student 7 : **7.25** <br>
Total score : 50.5 <br>

### ตัวอย่างการทำงาน: 

Number of students : **-1** <br>
Incorrect Input. Number of students : **7** <br>
Student 1 : **10** <br>
Student 2 : **6.5** <br>
Student 3 : **0** <br>
Student 4 : **9** <br>
Student 5 : **8.25** <br>
Student 6 : **9.50** <br>
Student 7 : **7.25** <br>
Total score : 50.5 <br>

<!-- grade_id: score, points: 0, grade: false, solution: true -->
```python
def run():
    s = input("Enter username :")
    i =0 
    suma = 0
    while True:
        if int(s) < 0 :
            s =float(input("Incorrect Input. Number of students:"))
        if i <= int(s)-1 :
            i+=1
            score = float(input("Student :"))
            suma += score
        if i==int(s):
            break
            
            
            
    print("Total score :",float(suma))
                     
        
```

```python
'''ตรวจสอบผลลัพธ์ของโปรแกรมด้วยการรันที่ cell นี้'''
run()
```

Output:
```
Enter username :0
```

<!-- grade_id: partial, points: 1, grade: true, solution: false -->
```python
'''ตรวจสอบผลการคำนวณ'''
from unittest.mock import patch
from io import StringIO
import math

def test(inputs,expected_output,exact_output): 
    with patch('builtins.input', side_effect=inputs), patch('sys.stdout', new_callable=StringIO) as mock_stdout:
        run()
        assert mock_stdout.getvalue() == expected_output,mock_stdout.getvalue()
```

<!-- grade_id: cell-5646d4ae088951fe, points: 1, grade: true, solution: false -->
```python
'''ตรวจสอบผลการคำนวณและการแสดงผลลัพธ์ให้เหมือนกับโจทย์ที่กำหนดให้'''
'''Number of students is 3'''
test(['3','10','6.5','0'],'Total score : 16.5\n',exact_output=True) 
test(['3','7','4','9'],'Total score : 20.0\n',exact_output=True) 
test(['3','6','5.5','9.8'],'Total score : 21.3\n',exact_output=True) 
test(['3','1.5','4.5','6.5'],'Total score : 12.5\n',exact_output=True) 
```

<!-- grade_id: exact, points: 1, grade: true, solution: false -->
```python
'''ตรวจสอบผลการคำนวณและการแสดงผลลัพธ์ให้เหมือนกับโจทย์ที่กำหนดให้'''
'''Number of students is 5'''
test(['5','10','6.5','0','9','8.25'],'Total score : 33.75\n',exact_output=True) 
test(['5','6','8','9','4','10'],'Total score : 37.0\n',exact_output=True) 
test(['5','4.25','6.25','1.25','8.75','9.25'],'Total score : 29.75\n',exact_output=True) 
test(['5','4.25','1','0','8.75','9.25'],'Total score : 23.25\n',exact_output=True) 
```

<!-- grade_id: cell-cbc47179f27fefee, points: 1, grade: true, solution: false -->
```python
'''ตรวจสอบผลการคำนวณและการแสดงผลลัพธ์ให้เหมือนกับโจทย์ที่กำหนดให้'''
'''Number of students is 7'''
test(['7','10','6.5','0','9','8.25','9.50','7.25'],'Total score : 50.5\n',exact_output=True) 
test(['7','4.25','6.25','1.25','8.75','9.25','7.25','9.75'],'Total score : 46.75\n',exact_output=True) 
test(['7','6','8','9','4','10','3','4'],'Total score : 44.0\n',exact_output=True) 
```

<!-- grade_id: cell-4234104d92125bb2, points: 1, grade: true, solution: false -->
```python
'''ตรวจสอบเคส inputs ติด ลบ'''
test(['-1','-5','7','10','6.5','0','9','8.25','9.50','7.25'],'Total score : 50.5\n',exact_output=True) 
test(['-5','7','10','6.5','0','9','8.25','9.50','7.25'],'Total score : 50.5\n',exact_output=True)
```

<!-- grade_id: cell-29679434cc1cbadb, points: 1, grade: true, solution: false -->
```python
'''ตรวจสอบผลการคำนวณและการแสดงผลลัพธ์ให้เหมือนกับโจทย์ที่กำหนดให้'''
test(['0'],'Total score : 0.0\n',exact_output=True)
```

<!-- grade_id: cell-714b60293955bf54, points: 1, grade: true, solution: false -->
```python
'''เทส ลับ'''
```