## Q1 (5 points)
จงเขียนโปรแกรมที่วนรับชื่อผู้ใช้จนกว่าจะเป็น Lisa

### ตัวอย่างการทำงาน: 

Enter username : **Lisa** <br>
Hello, Lisa

### ตัวอย่างการทำงาน: 

Enter username : **test** <br>
Incorrect. Enter again : **Lisa** <br>
Hello, Lisa

### ตัวอย่างการทำงาน: 
Enter username : **test** <br>
Incorrect. Enter again : **12345** <br>
Incorrect. Enter again : **Lisa** <br>
Hello, Lisa

### ตัวอย่างการทำงาน: 

Enter username : **test** <br>
Incorrect. Enter again : **1234** <br>
Incorrect. Enter again : **Jackson** <br>
Incorrect. Enter again : **lisa** <br>
Incorrect. Enter again : **Lisa** <br>
Hello, Lisa

<!-- grade_id: name, points: 0, grade: false, solution: true -->
```python
def run():
    s = input("Enter username :")
    while True:   
        if s != "Lisa":
            s =input("incorrect, Enter again : ")
        else:
            print("Hello,",s)
            break
```

```python
'''ตรวจสอบผลลัพธ์ของโปรแกรมด้วยการรันที่ cell นี้'''
run()
```

<!-- grade_id: partial, points: 0, grade: true, solution: false -->
```python
'''ตรวจสอบผลการคำนวณ'''
from unittest.mock import patch
from io import StringIO
import math

def test(inputs,expected_output,exact_output): 
    with patch('builtins.input', side_effect=inputs), patch('sys.stdout', new_callable=StringIO) as mock_stdout:
        run()
        assert mock_stdout.getvalue() == expected_output,mock_stdout.getvalue()
```

<!-- grade_id: exact, points: 1, grade: true, solution: false -->
```python
'''ตรวจสอบผลการคำนวณและการแสดงผลลัพธ์ให้เหมือนกับโจทย์ที่กำหนดให้'''
test(['Lisa'],'Hello, Lisa\n',exact_output=True) 
test(['test','Lisa'],'Hello, Lisa\n',exact_output=True) 
test(['Blackpink','Lisa'],'Hello, Lisa\n',exact_output=True)  
```

<!-- grade_id: cell-cbc47179f27fefee, points: 1, grade: true, solution: false -->
```python
'''ตรวจสอบผลการคำนวณและการแสดงผลลัพธ์ให้เหมือนกับโจทย์ที่กำหนดให้'''
test(['test','12345','Lisa'],'Hello, Lisa\n',exact_output=True)
test(['ABC','XYZ','Lisa'],'Hello, Lisa\n',exact_output=True)
test(['Jackson','Mario','Lisa'],'Hello, Lisa\n',exact_output=True)
```

<!-- grade_id: cell-4c547bf9d492d10f, points: 1, grade: true, solution: false -->
```python
'''ตรวจสอบผลการคำนวณและการแสดงผลลัพธ์ให้เหมือนกับโจทย์ที่กำหนดให้'''
test(['test','12345','xxx','Lisa'],'Hello, Lisa\n',exact_output=True)
test(['Shark','Monkey','Chicken','Lisa'],'Hello, Lisa\n',exact_output=True)
test(['Pee','Prim','Tee','Lisa'],'Hello, Lisa\n',exact_output=True)
```

<!-- grade_id: cell-0518eb220bac84ec, points: 1, grade: true, solution: false -->
```python
'''ตรวจสอบผลการคำนวณและการแสดงผลลัพธ์ให้เหมือนกับโจทย์ที่กำหนดให้'''
test(['test','12345','Jackson','lisa','Lisa'],'Hello, Lisa\n',exact_output=True)
test(['Bambam','Jennie','Jackson','Jisoo','Lisa'],'Hello, Lisa\n',exact_output=True)
test(['Oat','Pop','Nont','Mario','Lisa'],'Hello, Lisa\n',exact_output=True)
```

<!-- grade_id: cell-a1e07cb980135999, points: 1, grade: true, solution: false -->
```python
'''เทส ลับ'''
```