---
## Q2 (5 points)
จงเขียนโปรแกรมที่รับน้ำหนัก (kg.) และส่วนสูง (m.) จากผู้ใช้มาคำนวณ bmi จาก น้ำหนัก (kg.)/ส่วนสูง
(m.) ยกกำลัง 2 แล้วแสดงผลลัพธ์ว่าผู้ใช้มีรูปร่างในระดับใดดังนี้

ต่ำกว่า 18.5 ผอม

ไม่น้อยกว่า 18.5 แต่ต่ำกว่า 23.0 รูปร่างปกติ

ไม่น้อยกว่า 23.0 แต่ต่ำกว่า 25.0 รูปร่างอ้วน

ไม่น้อยกว่า 25.0 แต่ต่ำกว่า 30.0 อ้วนระดับ 1

ไม่น้อยกว่า 30.0 อ้วนระดับ 2

### ตัวอย่างการทำงาน: 

Weight (kg.) : **61.3** <br>
Height (m.) : **1.65** <br>
รูปร่างปกติ

### ตัวอย่างการทำงาน: 

Weight (kg.) : **50.3** <br>
Height (m.) : **1.67**<br>
ผอม

### ตัวอย่างการทำงาน:

Weight (kg.) : **72.3** <br> Height (m.) : 1.5<br>อ้วนระดับ 2


<!-- grade_id: bmi, points: 0, grade: false, solution: true -->
```python
def run():
    a =float(input())
    b = float(input())
    c= a/ (b**2)
    if c < 18.5  :
        print("ผอม")
    elif c >= 18.5 and c <23:
        print("รูปร่างปกติ")
    elif c>= 23 and  c<25:
        print("รูปร่างอ้วน")
    elif c>=25 and c< 30:
        print("อ้วนระดับ 1")
    else :
        print("อ้วนระดับ 2")
    
```

```python
'''ตรวจสอบผลลัพธ์ของโปรแกรมด้วยการรันที่ cell นี้'''
run()
```

Output:
```
75
1.75
24.489795918367346
รูปร่างอ้วน
```

<!-- grade_id: partial, points: 0, grade: true, solution: false -->
```python
'''ตรวจสอบผลการคำนวณ'''
from unittest.mock import patch
from io import StringIO
import math

def test(inputs,expected_output,exact_output): 
    with patch('builtins.input', side_effect=inputs), patch('sys.stdout', new_callable=StringIO) as mock_stdout:
        run()
        if exact_output:
            assert mock_stdout.getvalue() == expected_output,mock_stdout.getvalue()
        else:
            assert expected_output in mock_stdout.getvalue(),mock_stdout.getvalue()
```

<!-- grade_id: exact, points: 1, grade: true, solution: false -->
```python
'''ตรวจสอบผลการคำนวณและการแสดงผลลัพธ์ให้เหมือนกับโจทย์ที่กำหนดให้'''
test(['61.3','1.65'],'รูปร่างปกติ\n',exact_output=True) 
test(['74','2.0'],'รูปร่างปกติ\n',exact_output=True)   
```

Output:
```
Input In [11], in <cell line: 3>()
      1 '''ตรวจสอบผลการคำนวณและการแสดงผลลัพธ์ให้เหมือนกับโจทย์ที่กำหนดให้'''
      2 test(['61.3','1.65'],'รูปร่างปกติ\n',exact_output=True) 
----> 3 test(['74','2.0'],'รูปร่างปกติ\n',exact_output=True)

Input In [10], in test(inputs, expected_output, exact_output)
      8 run()
      9 if exact_output:
---> 10     assert mock_stdout.getvalue() == expected_output,mock_stdout.getvalue()
     11 else:
     12     assert expected_output in mock_stdout.getvalue(),mock_stdout.getvalue()

AssertionError: ผอม
```

<!-- grade_id: cell-015b162bcc29603f, points: 1, grade: true, solution: false -->
```python
'''ตรวจสอบผลการคำนวณและการแสดงผลลัพธ์ให้เหมือนกับโจทย์ที่กำหนดให้'''
test(['50.3','1.67'],'ผอม\n',exact_output=True) 
```

<!-- grade_id: cell-b40afe32dd830024, points: 1, grade: true, solution: false -->
```python
'''ตรวจสอบผลการคำนวณและการแสดงผลลัพธ์ให้เหมือนกับโจทย์ที่กำหนดให้'''
test(['72.3','1.5'],'อ้วนระดับ 2\n',exact_output=True)
```

<!-- grade_id: cell-8ae0f1eee35b9354, points: 1, grade: true, solution: false -->
```python
'''ตรวจสอบผลการคำนวณและการแสดงผลลัพธ์ให้เหมือนกับโจทย์ที่กำหนดให้'''
test(['73','1.68'],'อ้วนระดับ 1\n',exact_output=True) 
test(['100','2.0'],'อ้วนระดับ 1\n',exact_output=True)
```

<!-- grade_id: cell-0b1407a34c6f1456, points: 1, grade: true, solution: false -->
```python
'''ตรวจสอบผลการคำนวณและการแสดงผลลัพธ์ให้เหมือนกับโจทย์ที่กำหนดให้'''
test(['75','1.75'],'รูปร่างอ้วน\n',exact_output=True) 
test(['92','2.0'],'รูปร่างอ้วน\n',exact_output=True)
```