## Q3 (4 points)
จงเขียนโปรแกรมที่วนรับอุณหภูมิสูงสุดของแต่ละวันในหนึ่งสัปดาห์ แล้วแสดงว่าวันใดมีอุณหภูมิลดลงจากวันก่อนหน้า

### ตัวอย่างการทำงาน:

Day 1 : **32.6** <br>
Day 2 : **34.4** <br>
Day 3 : **35.0** <br>
Day 4 : **34.6** <br>
Temperature dropped on day 4 <br>
Day 5 : **35.0** <br>
Day 6 : **35.6** <br>
Day 7 : **36** <br>

### ตัวอย่างการทำงาน: 

Day 1 : **32.6** <br>
Day 2 : **34.4** <br>
Day 3 : **35.0** <br>
Day 4 : **34.6** <br>
Temperature dropped on day 4 <br>
Day 5 : **35.2** <br>
Day 6 : **35.6** <br>
Day 7 : **34** <br>
Temperature dropped on day 7

### ตัวอย่างการทำงาน: 

Day 1 : **32.6** <br>
Day 2 : **34.4** <br>
Day 3 : **35.0** <br>
Day 4 : **34.6** <br>
Temperature dropped on day 4 <br>
Day 5 : **34.0** <br>
Temperature dropped on day 5 <br>
Day 6 : **35.6** <br>
Day 7 : **34** <br>
Temperature dropped on day 7

<!-- grade_id: temp, points: 0, grade: false, solution: true -->
```python
def run():
    i=0
    t =0
    while i <= 6:
        x = float(input())
        i += 1
        if x - t >= 0:
            t = x
        else:
            print("Temperature dropped on day "+str(i))
            t = x
        
        
```

```python
'''ตรวจสอบผลลัพธ์ของโปรแกรมด้วยการรันที่ cell นี้'''
run()
```

Output:
```
1
2
3
4
5
6
7
```

<!-- grade_id: partial, points: 0, grade: true, solution: false -->
```python
'''ตรวจสอบผลการคำนวณ'''
from unittest.mock import patch
from io import StringIO
import math

def test(inputs,expected_output,exact_output): 
    with patch('builtins.input', side_effect=inputs), patch('sys.stdout', new_callable=StringIO) as mock_stdout:
        run()
        assert mock_stdout.getvalue() == expected_output,mock_stdout.getvalue()
```

<!-- grade_id: cell-cbc47179f27fefee, points: 1, grade: true, solution: false -->
```python
'''ตรวจสอบผลการคำนวณและการแสดงผลลัพธ์ให้เหมือนกับโจทย์ที่กำหนดให้'''
'''Drop 1 day'''
test(['32.6','34.4','35.0','34.6','35','35.6','36'],'Temperature dropped on day 4\n',exact_output=True) 
test(['32.6','31.4','35.0','35.6','35.8','35.9','36'],'Temperature dropped on day 2\n',exact_output=True) 
test(['32.6','33.4','35.0','35.6','35.8','35.9','34'],'Temperature dropped on day 7\n',exact_output=True) 
```

<!-- grade_id: exact, points: 1, grade: true, solution: false -->
```python
'''ตรวจสอบผลการคำนวณและการแสดงผลลัพธ์ให้เหมือนกับโจทย์ที่กำหนดให้'''
'''Drop 2 days'''
test(['32.6','34.4','35.0','34.6','35.2','35.6','34'],'Temperature dropped on day 4\nTemperature dropped on day 7\n',exact_output=True) 
test(['33.6','34.4','32.0','35.6','37.2','31.6','35'],'Temperature dropped on day 3\nTemperature dropped on day 6\n',exact_output=True) 
test(['43.6','40.4','42.0','45.6','30.2','31.6','35'],'Temperature dropped on day 2\nTemperature dropped on day 5\n',exact_output=True) 
```

<!-- grade_id: cell-5646d4ae088951fe, points: 1, grade: true, solution: false -->
```python
'''ตรวจสอบผลการคำนวณและการแสดงผลลัพธ์ให้เหมือนกับโจทย์ที่กำหนดให้'''
'''Drop 3 days'''
test(['32.6','34.4','35.0','34.6','34','35.6','34'],'Temperature dropped on day 4\nTemperature dropped on day 5\nTemperature dropped on day 7\n',exact_output=True) 
test(['43.6','40.4','39.0','45.6','30.2','31.6','35'],'Temperature dropped on day 2\nTemperature dropped on day 3\nTemperature dropped on day 5\n',exact_output=True)
test(['33.6','34.4','32.0','31.5','37.2','31.6','35'],'Temperature dropped on day 3\nTemperature dropped on day 4\nTemperature dropped on day 6\n',exact_output=True) 
test(['45','45','45','45','45','45','45'],'',exact_output=True)
```

<!-- grade_id: cell-4cf323a44b865e68, points: 1, grade: true, solution: false -->
```python
'''เทสลับ'''
```