## Q2 (4 points)
จงเขียนโปรแกรมที่รับจำนวนเต็มบวก n แล้วหาจำนวนเต็มบวกที่หาร n ลงตัว

### ตัวอย่างการทำงาน: 

Enter an integer : **1** <br>
1 is divisible by:<br>
1<br>

### ตัวอย่างการทำงาน: 

Enter an integer : **12** <br>
12 is divisible by: <br>
1 <br>
2 <br>
3 <br>
4 <br>
6 <br>
12 <br>

### ตัวอย่างการทำงาน:

Enter an integer : **65** <br>
65 is divisible by: <br>
1 <br>
5 <br>
13 <br>
65 <br>

### ตัวอย่างการทำงาน: 

Enter an integer : **211** <br>
211 is divisible by: <br>
1 <br>
211 <br>

<!-- grade_id: division, points: 0, grade: false, solution: true -->
```python
def run():
    x = int(input())
    print(str(x),"is divisible by:")
    for i in range(1,x+1):
        if x%i ==0:
            print(int(i))
```

```python
'''ตรวจสอบผลลัพธ์ของโปรแกรมด้วยการรันที่ cell นี้'''
run()
```

Output:
```
211
1
211
```

<!-- grade_id: partial, points: 0, grade: true, solution: false -->
```python
'''ตรวจสอบผลการคำนวณ'''
from unittest.mock import patch
from io import StringIO
import math

def test(inputs,expected_output,exact_output): 
    with patch('builtins.input', side_effect=inputs), patch('sys.stdout', new_callable=StringIO) as mock_stdout:
        run()
        assert mock_stdout.getvalue() == expected_output,mock_stdout.getvalue()
```

<!-- grade_id: exact, points: 1, grade: true, solution: false -->
```python
'''ตรวจสอบผลการคำนวณและการแสดงผลลัพธ์ให้เหมือนกับโจทย์ที่กำหนดให้'''
'''Even numbers'''
test(['12'],'12 is divisible by:\n1\n2\n3\n4\n6\n12\n',exact_output=True) 
test(['128'],'128 is divisible by:\n1\n2\n4\n8\n16\n32\n64\n128\n',exact_output=True) 
test(['1024'],'1024 is divisible by:\n1\n2\n4\n8\n16\n32\n64\n128\n256\n512\n1024\n',exact_output=True) 
```

<!-- grade_id: cell-a07ef84aa3313800, points: 1, grade: true, solution: false -->
```python
'''ตรวจสอบผลการคำนวณและการแสดงผลลัพธ์ให้เหมือนกับโจทย์ที่กำหนดให้'''
'''Odd numbers'''
test(['1'],'1 is divisible by:\n1\n',exact_output=True) 
test(['65'],'65 is divisible by:\n1\n5\n13\n65\n',exact_output=True) 
test(['111'],'111 is divisible by:\n1\n3\n37\n111\n',exact_output=True) 
test(['1537'],'1537 is divisible by:\n1\n29\n53\n1537\n',exact_output=True) 
```

<!-- grade_id: cell-ad91ccd55eeb0c2f, points: 1, grade: true, solution: false -->
```python
'''ตรวจสอบผลการคำนวณและการแสดงผลลัพธ์ให้เหมือนกับโจทย์ที่กำหนดให้'''
'''Prime numbers'''
test(['211'],'211 is divisible by:\n1\n211\n',exact_output=True)
test(['557'],'557 is divisible by:\n1\n557\n',exact_output=True)
test(['941'],'941 is divisible by:\n1\n941\n',exact_output=True)
```

<!-- grade_id: cell-9d4df4635fcf41eb, points: 1, grade: true, solution: false -->
```python
'''เทสลับ'''
```