## Q4 (4 points)
จงเขียนโปรแกรมที่วนรับจำนวนเงินที่ต้องการเบิกจากบัญชี โดยจะยอมให้เบิกถ้ายังมีเงินในบัญชีพอ และวนรับจำนวนเงินที่ต้องการเบิกไปจนเงินในบัญชีหมดพอดี (กำหนดให้เริ่มต้นมีเงินในบัญชี 50,000 บาท)


### ตัวอย่างการทำงาน: 

withdraw : **50000** <br>
Balance is 0

### ตัวอย่างการทำงาน: 

withdraw : **45000**  <br>
withdraw : **5000**  <br>
Balance is 0

### ตัวอย่างการทำงาน: 

withdraw : **60000**  <br>
Insufficient fund. <br>
withdraw : **40000**  <br>
withdraw : **20000**  <br>
Insufficient fund.<br>
withdraw : **10000**  <br>
Balance is 0

<!-- grade_id: withdraw, points: 0, grade: false, solution: true -->
```python
def run():
    Ba = 50000
    while Ba > 0:
        x = int(input())
        if x > Ba:
            print("Insufficient fund.")
        else:
            Ba = Ba - x
    print("Balance is 0")
        
```

```python
'''ตรวจสอบผลลัพธ์ของโปรแกรมด้วยการรันที่ cell นี้'''
run()
```

Output:
```
4000
46000
50000
Insufficient fund.
20000
26000
```

<!-- grade_id: partial, points: 0, grade: true, solution: false -->
```python
'''ตรวจสอบผลการคำนวณ'''
from unittest.mock import patch
from io import StringIO
import math

def test(inputs,expected_output,exact_output): 
    with patch('builtins.input', side_effect=inputs), patch('sys.stdout', new_callable=StringIO) as mock_stdout:
        run()
        assert mock_stdout.getvalue() == expected_output,mock_stdout.getvalue()
```

<!-- grade_id: exact, points: 1, grade: true, solution: false -->
```python
'''ตรวจสอบผลการคำนวณและการแสดงผลลัพธ์ให้เหมือนกับโจทย์ที่กำหนดให้'''
test(['50000'],'Balance is 0\n',exact_output=True) 
test(['25000','0','25000'],'Balance is 0\n',exact_output=True) 
```

<!-- grade_id: cell-5646d4ae088951fe, points: 1, grade: true, solution: false -->
```python
'''ตรวจสอบผลการคำนวณและการแสดงผลลัพธ์ให้เหมือนกับโจทย์ที่กำหนดให้'''
test(['45000','5000'],'Balance is 0\n',exact_output=True)  
test(['20000','30000'],'Balance is 0\n',exact_output=True)  
test(['98723','50000'],'Insufficient fund.\nBalance is 0\n',exact_output=True)  
test(['73568','50000'],'Insufficient fund.\nBalance is 0\n',exact_output=True) 
```

<!-- grade_id: cell-cbc47179f27fefee, points: 1, grade: true, solution: false -->
```python
'''ตรวจสอบผลการคำนวณและการแสดงผลลัพธ์ให้เหมือนกับโจทย์ที่กำหนดให้'''
test(['60000','40000','20000','10000'],'Insufficient fund.\nInsufficient fund.\nBalance is 0\n',exact_output=True)
test(['100000','60000','30000','20000'],'Insufficient fund.\nInsufficient fund.\nBalance is 0\n',exact_output=True)
test(['98765','56789','43210','6790'],'Insufficient fund.\nInsufficient fund.\nBalance is 0\n',exact_output=True)
test(['98765','43210','12345','6660','130'],'Insufficient fund.\nInsufficient fund.\nBalance is 0\n',exact_output=True)
```

<!-- grade_id: cell-d4781a6782b18bd3, points: 1, grade: true, solution: false -->
```python
'''เทสลับ'''
```