## Q1 (4 points)
จงเขียนโปรแกรมที่รับคำมาต่อกันเป็นประโยค โดยจะหยุดรับเมื่อรับ . มา

### ตัวอย่างการทำงาน: 

Next word : **.** <br>
Sentence:

### ตัวอย่างการทำงาน: 

Next word : **Bye** <br>
Next word : **bye** <br>
Next word : **.** <br>
Sentence: Bye bye

### ตัวอย่างการทำงาน: 

Next word : **My** <br>
Next word : **name** <br>
Next word : **is** <br>
Next word : **next** <br>
Next word : **word** <br>
Next word : **.** <br>
Sentence: My name is next word

### ตัวอย่างการทำงาน: 

Next word : **My** <br>
Next word : **name** <br>
Next word : **is** <br>
Next word : **Chicky** <br>
Next word : **Chicky!** <br>
Next word : **.** <br>
Sentence: My name is Chicky Chicky!

<!-- grade_id: name, points: 0, grade: false, solution: true -->
```python
def run():
    x = []
    n = input()
    dot = True
    while n != "." :
        x.append(n)
        n = input()
        dot = False
    if dot == True :
        print("Sentence:")
    else :
        print("Sentence: "+ " ".join(x))
        
            
```

Output:
```
Input In [43]
    if x==[.]:
           ^
SyntaxError: invalid syntax
```

```python
'''ตรวจสอบผลลัพธ์ของโปรแกรมด้วยการรันที่ cell นี้'''
run()
```

Output:
```
k
u
y
.
Sentence:
```

<!-- grade_id: partial, points: 0, grade: true, solution: false -->
```python
'''ตรวจสอบผลการคำนวณ'''
from unittest.mock import patch
from io import StringIO
import math

def test(inputs,expected_output,exact_output): 
    with patch('builtins.input', side_effect=inputs), patch('sys.stdout', new_callable=StringIO) as mock_stdout:
        run()
        assert mock_stdout.getvalue() == expected_output,mock_stdout.getvalue()
```

<!-- grade_id: exact, points: 1, grade: true, solution: false -->
```python
'''ตรวจสอบผลการคำนวณและการแสดงผลลัพธ์ให้เหมือนกับโจทย์ที่กำหนดให้'''
test(['.'],'Sentence:\n',exact_output=True) 
test(['Welcome','.'],'Sentence: Welcome\n',exact_output=True)
test(['Comprog','.'],'Sentence: Comprog\n',exact_output=True)
```

<!-- grade_id: cell-cbc47179f27fefee, points: 1, grade: true, solution: false -->
```python
'''ตรวจสอบผลการคำนวณและการแสดงผลลัพธ์ให้เหมือนกับโจทย์ที่กำหนดให้'''
test(['Bye','bye','.'],'Sentence: Bye bye\n',exact_output=True) 
test(['I','love','Comprog','.'],'Sentence: I love Comprog\n',exact_output=True) 
```

<!-- grade_id: cell-4c547bf9d492d10f, points: 1, grade: true, solution: false -->
```python
'''ตรวจสอบผลการคำนวณและการแสดงผลลัพธ์ให้เหมือนกับโจทย์ที่กำหนดให้'''
test(['My','name','is','next','word','.'],'Sentence: My name is next word\n',exact_output=True)  
test(['My','name','is','Chicky','Chicky!','.'],'Sentence: My name is Chicky Chicky!\n',exact_output=True)
```

<!-- grade_id: cell-63c7bc3d8547f871, points: 1, grade: true, solution: false -->
```python
'''เทสลับ'''
```