## Q2 (3 points)
จงเขียนโปรแกรมที่รับตำแหน่ง (x,y) เริ่มต้นของหุ่นยนต์จากผู้ใช้ (x และ y เป็นจำนวนเต็ม) <br>
แล้วอ่านลำดับของคำสั่งที่ให้หุ่นยนต์เดินไปทางซ้าย (L) ขวา (R) ขึ้นบน (U) ลงล่าง (D) ครั้งละ 1 หน่วย <br>
จากไฟล์ **move\<n\>.txt** แล้วหาว่าเมื่อจบโปรแกรมหุ่นยนต์จะไปอยู่ที่ตำแหน่งใด โดยให้มีกำแพงที่เส้น **x=10, x=-10, y=10 และ y=-10**<br>
ถ้าการเดินของหุ่นยนต์จะทำให้หุ่นยนต์ชนกำแพง หุ่นยนต์จะไม่เคลื่อนที่จากตำแหน่งเดิม เช่น 
เมื่อหุ่นยนต์อยู่ที่ (9,2) และได้คำสั่ง R หุ่นจะชนกำแพงที่เส้น x=10 ดังนั้นจะไม่ขยับตำแหน่ง
**ผู้ใช้จะพิมพ์ตำแหน่งเริ่มต้นของหุ่นยนต์ที่อยู่ในกล่องเท่านั้น**

### ตัวอย่างการทำงาน: 

ไฟล์ move3.txt: <br>
L <br>
L <br>
D <br>
R <br>
R <br>
R <br>
R <br>
----
Choose your movefile: **move3.txt** <br>
Initial position : **9,-9** <br>
Robot stops at 9,-9

### ตัวอย่างการทำงาน: 

ไฟล์ move3.txt:<br>
L <br>
L <br>
D <br>
R <br>
R <br>
R <br>
R <br>
---
Choose your movefile: **move3.txt** <br>
Initial position : **8,4** <br>
Robot stops at 9,3

### ตัวอย่างการทำงาน:

ไฟล์ move4.txt: <br>
L <br>
L <br>
X <br>
U <br>
D <br>
Z <br>
R <br>
----
Choose your movefile: **move4.txt** <br>
Initial position : **1,2** <br>
Invalid command

<!-- grade_id: answer, points: 0, grade: false, solution: true -->
```python
def run():
    choose = input("Choose your movefile: ")
    if "move" not in choose:
        return
    x,y = input("Initial position: ").split(',')
    x,y = int(x),int(y)
    if x>10 or y>10 or x<-10 or y<-10:
        print("Invalid command")
        return
    f = open(choose,'r')
    for i in f:
        if "L" in i:
            if x!= -9:
                x-=1
        elif "R" in i:
            if x != 9:
                x+=1
        elif "U" in i:
            if y != 9:
                y+=1
        elif "D" in i:
            if y != -9:
                y-=1
        else:
            print("Invalid command")
            return
    print("Robot stops at "+str(x)+","+str(y))
```

```python
'''ตรวจสอบผลลัพธ์ของโปรแกรมด้วยการรันที่ cell นี้'''
run()
```

Output:
```
Choose your movefile: move4.txt
Initial position: 1,1
L

L

D

R

R

R

R
Robot stops at 3,0
```

<!-- grade_id: partial, points: 0, grade: true, solution: false -->
```python
'''ตรวจสอบผลการคำนวณ'''
from unittest.mock import patch
from io import StringIO
import math

def test(inputs,expected_output,exact_output): 
    with patch('builtins.input', side_effect=inputs), patch('sys.stdout', new_callable=StringIO) as mock_stdout:
        run()
        assert mock_stdout.getvalue() == expected_output,mock_stdout.getvalue()
```

<!-- grade_id: exact, points: 1, grade: true, solution: false -->
```python
'''ตรวจสอบผลการคำนวณและการแสดงผลลัพธ์ให้เหมือนกับโจทย์ที่กำหนดให้'''
'''Robot can stop at position x,y'''
test(['move3.txt','9,-9'],'Robot stops at 9,-9\n',exact_output=True) #ex
test(['move3.txt','8,4'],'Robot stops at 9,3\n',exact_output=True) #ex
test(['move3.txt','9,5'],'Robot stops at 9,4\n',exact_output=True) # ++
test(['move3.txt','-5,5'],'Robot stops at -3,4\n',exact_output=True)  #-+
test(['move3.txt','-1,-1'],'Robot stops at 1,-2\n',exact_output=True)  #--
```

<!-- grade_id: cell-a07ef84aa3313800, points: 1, grade: true, solution: false -->
```python
'''ตรวจสอบผลการคำนวณและการแสดงผลลัพธ์ให้เหมือนกับโจทย์ที่กำหนดให้'''
'''Invalid command'''
test(['move4.txt','15,20'],'Invalid command\n',exact_output=True) #
test(['move4.txt','1,1'],'Invalid command\n',exact_output=True) #
test(['move4.txt','1,2'],'Invalid command\n',exact_output=True) #ex
```

<!-- grade_id: cell-ad91ccd55eeb0c2f, points: 1, grade: true, solution: false -->
```python
'''เทส ลับ'''
```