## Q3 (3 points)
จงเขียนฟังก์ชัน f(x) ที่คำนวณ 1/x! และฟังก์ชัน summation(a,b) ที่ใช้ฟังก์ชัน f และ คำนวณผลรวมของ f(x) เมื่อ x = a, a+1,..., b (นั่นคือ 1/a! + 1/(a+1)! + ... + 1/b!) โดยที่ a ต้องไม่มากกว่า b <br>
โดยใช้ ฟังก์ชัน factorial จาก module math <br>
 
จากนั้นให้เขียนโปรแกรมหลักที่รับ**จำนวนเต็มบวก 2 จำนวน m และ n** (m อาจมากกว่าหรือน้อยกว่า n) แล้วเรียกใช้ฟังก์ชัน summation เพื่อ คำนวณหาผลรวมของ f(x) เมื่อ x มีค่าตั้งแต่ m ถึง n <br>
**โดยแสดงผลลัพธ์เป็นตัวเลขทศนิยม 4 ตำแหน่ง**

### ตัวอย่างการทำงาน: 

Enter m: **1** <br>
Enter n: **3** <br>
1.6667


### ตัวอย่างการทำงาน: 

Enter m: **3** <br>
Enter n: **1** <br>
1.6667

### ตัวอย่างการทำงาน: 

Enter m: **3** <br>
Enter n: **6** <br>
0.2181

<!-- grade_id: answer, points: 0, grade: false, solution: true -->
```python
import math
def f(x):
    return (1/math.factorial(x))
def summation(a,b):
    x = 0
    while a<=b:
        x += f(a)
        a +=1
    print(round(x,4))
def run():
    m = int(input("m : "))
    n = int(input("n : "))
    if m<=n:
        return summation(m,n)
    else:
        return summation(n,m)
```

```python
'''ตรวจสอบผลลัพธ์ของโปรแกรมด้วยการรันที่ cell นี้'''
run()
```

Output:
```
m : 3
n : 6
0.2181
```

<!-- grade_id: partial, points: 0, grade: true, solution: false -->
```python
'''ตรวจสอบผลการคำนวณ'''
from unittest.mock import patch
from io import StringIO
import math

def test(inputs,expected_output,exact_output): 
    with patch('builtins.input', side_effect=inputs), patch('sys.stdout', new_callable=StringIO) as mock_stdout:
        run()
        assert mock_stdout.getvalue() == expected_output,mock_stdout.getvalue()
```

<!-- grade_id: cell-4f7b8ee5a68b76a7, points: 1, grade: true, solution: false -->
```python
'''ตรวจสอบผลการคำนวณและการแสดงผลลัพธ์ให้เหมือนกับโจทย์ที่กำหนดให้'''
test(['1','3'],'1.6667\n',exact_output=True)  
test(['3','6'],'0.2181\n',exact_output=True) 
test(['5','6'],'0.0097\n',exact_output=True)
test(['3','4'],'0.2083\n',exact_output=True)
```

<!-- grade_id: cell-9becccc26c51e026, points: 1, grade: true, solution: false -->
```python
'''ตรวจสอบผลการคำนวณและการแสดงผลลัพธ์ให้เหมือนกับโจทย์ที่กำหนดให้'''
test(['3','1'],'1.6667\n',exact_output=True) 
test(['9','2'],'0.7183\n',exact_output=True)
test(['5','2'],'0.7167\n',exact_output=True)
```

<!-- grade_id: exact, points: 1, grade: true, solution: false -->
```python
'''เทสลับ'''
```