## Q1 (3 points)
จงเขียนโปรแกรมที่รับตำแหน่ง (x,y) (x และ y เป็นจำนวนเต็ม) ซึ่งเป็นจุดเริ่มต้นของหุ่นยนต์จากผู้ใช้ แล้วอ่านลำดับของคำสั่งที่ให้หุ่นยนต์เดินไปทางซ้าย (L) ขวา (R) ขึ้นบน (U) ลงล่าง (D) ครั้งละ 1 หน่วย จากไฟล์ **move\<n\>.txt** แล้วหาว่าเมื่อจบโปรแกรมหุ่นยนต์จะไปอยู่ที่ตำแหน่งใดและแสดงผลลัพธ์ดังตัวอย่างข้างล่างนี้ <br>


### ตัวอย่างการทำงาน: 

ไฟล์ move1.txt:<br>
L <br>
L <br>
U <br>
U <br>
D <br>
R <br>
R <br>
U <br>
U <br>
L <br>
L <br>
----
Choose your movefile: **move1.txt** <br>
Initial position : **10,20** <br>
Robot stops at 8,23


### ตัวอย่างการทำงาน: 

ไฟล์ move2.txt: <br>
L <br>
L <br>
X <br>
U <br>
D <br>
Z <br>
R <br>
----
Choose your movefile: **move2.txt** <br>
Initial position : **10,20** <br>
Invalid command

<!-- grade_id: answer, points: 0, grade: false, solution: true -->
```python
def run():
    choose = input("Choose your movefile: ")
    if "move" not in choose:
        return
    x,y = input("Initial position: ").split(',')
    x,y = int(x),int(y)
    f = open(choose,'r')
    for i in f:
        if "L" in i:
            x-=1
        elif "R" in i:
            x+=1
        elif "U" in i:
            y+=1
        elif "D" in i:
            y-=1
            
        else:
            print("Invalid command")
            return
    print("Robot stops at "+str(x)+","+str(y))
```

```python
'''ตรวจสอบผลลัพธ์ของโปรแกรมด้วยการรันที่ cell นี้'''
run()
```

Output:
```
Choose your movefile: move
File ~/opt/anaconda3/lib/python3.9/site-packages/ipykernel/kernelbase.py:1075, in Kernel.raw_input(self, prompt)
   1071 if not self._allow_stdin:
   1072     raise StdinNotImplementedError(
   1073         "raw_input was called, but this frontend does not support input requests."
   1074     )
-> 1075 return self._input_request(
   1076     str(prompt),
   1077     self._parent_ident["shell"],
   1078     self.get_parent("shell"),
   1079     password=False,
   1080 )

File ~/opt/anaconda3/lib/python3.9/site-packages/ipykernel/kernelbase.py:1120, in Kernel._input_request(self, prompt, ident, parent, password)
   1117             break
   1118 except KeyboardInterrupt:
   1119     # re-raise KeyboardInterrupt, to truncate traceback
-> 1120     raise KeyboardInterrupt("Interrupted by user") from None
   1121 except Exception:
   1122     self.log.warning("Invalid Message:", exc_info=True)

KeyboardInterrupt: Interrupted by user
```

<!-- grade_id: partial, points: 0, grade: true, solution: false -->
```python
'''ตรวจสอบผลการคำนวณ'''
from unittest.mock import patch
from io import StringIO
import math

def test(inputs,expected_output,exact_output): 
    with patch('builtins.input', side_effect=inputs), patch('sys.stdout', new_callable=StringIO) as mock_stdout:
        run()
        assert mock_stdout.getvalue() == expected_output,mock_stdout.getvalue()
```

<!-- grade_id: exact, points: 1, grade: true, solution: false -->
```python
'''ตรวจสอบผลการคำนวณและการแสดงผลลัพธ์ให้เหมือนกับโจทย์ที่กำหนดให้'''
'''Robot can stop at position x,y'''
test(['move1.txt','10,20'],'Robot stops at 8,23\n',exact_output=True) # ++
test(['move1.txt','7,5'],'Robot stops at 5,8\n',exact_output=True) #++
test(['move1.txt','100,-20'],'Robot stops at 98,-17\n',exact_output=True)  #+-
test(['move1.txt','-100,50'],'Robot stops at -102,53\n',exact_output=True)  #-+
test(['move1.txt','-100,-100'],'Robot stops at -102,-97\n',exact_output=True)  #--
```

<!-- grade_id: cell-cbc47179f27fefee, points: 1, grade: true, solution: false -->
```python
'''ตรวจสอบผลการคำนวณและการแสดงผลลัพธ์ให้เหมือนกับโจทย์ที่กำหนดให้'''
'''Invalid command'''
test(['move2.txt','15,20'],'Invalid command\n',exact_output=True) #
test(['move2.txt','1,1'],'Invalid command\n',exact_output=True) #
```

<!-- grade_id: cell-4c547bf9d492d10f, points: 1, grade: true, solution: false -->
```python
'''เทส ลับ'''
```